```python
import pandas as pd
```

# Project 1 - Part 2

This project depends on a completed Part 1.  Make sure to complete Part 1 and save the CSV file (by running the last code cell) before you start on Part 2.

## Part 2a: Finding the Average Color of a Region of an Image

In Part 1, you found the average color of an image.  For this part, find the average color of a region of the image array `lab`, a starting location (`x` and `y`), and a `width` and `height`.  You need to return the same data type as in Part 1 (a dictionary with `L`, `a`, and `b`).

*(Hint: Your source code from Part 1 should help you out here!)*

<!-- grade_id: cell-8696f7ce092e5e9c, points: 0, grade: false, solution: true -->
```python
def findAverageColor(lab, start_x, start_y, width, height):
    # Your code here
    ...

    
```

<!-- grade_id: cell-d0f665ec48a6615e, points: 1, grade: true, solution: false -->
```python
## == *CORRECTED* TEST CASES for Part 2a ==
# - This read-only cell contains test cases for your previous cell(s).
# - If this cell runs without any errors in the output, you PASSED all test cases!
# - If this cell results errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
info = '\N{INFORMATION SOURCE}'

%run project1-settings.ipynb

def verify(result, L, a, b, info=None):
    import math
    
    assert( math.isclose( result['L'], L, rel_tol=0.01) ), f'Your `L` value is incorrect. {info}'
    assert( math.isclose( result['a'], a, rel_tol=0.01 ) ), f'Your `a` value is incorrect. {info}'
    assert( math.isclose( result['b'], b, rel_tol=0.01 ) ), f'Your `b` value is incorrect. {info}'


lab1 = fileToLabArray("test.png", resize=False)
test = findAverageColor(lab1, 0, 0, 4, 4)


# Tests for correct return type:
assert( test != None ), "Your findAverageColor function must return a value (right now it's returning nothing)."
assert( type(test) == type({}) ), f"Your findAverageColor function must return a dictionary (right now it's returning {type(test)})."

assert( 'L' in test ), "Your findAverageColor must return a value for 'L'."
assert( 'a' in test ), "Your findAverageColor must return a value for 'a'."
assert( 'b' in test ), "Your findAverageColor must return a value for 'b'."


# Lots of tests for correctness:
print(f"{info} Your test.png values (0, 0) -> (4, 4): (L={test['L']}, a={test['a']}, b={test['b']})")
verify( test, L=47.197, a=49.034, b=38.609 )


test2 = findAverageColor(lab1, 0, 0, 2, 2)
print(f"{info} Your test.png values (0, 0) -> (2, 2): (L={test2['L']}, a={test2['a']}, b={test2['b']})")
verify( test2, L=54.244, a=59.314, b=52.980, info="This test case tests only a 2x2 pixel area of the image.  Make sure you are looking at only a width x height region and not the full image.")


lab3 = fileToLabArray("test3.png", resize=False)
test3 = findAverageColor(lab3, 0, 0, 4, 6)
print(f"{info} Your test3.png values start=(0, 0) w/ width=4, height=6: (L={test3['L']}, a={test3['a']}, b={test3['b']})")
verify( test3, L=46.414, a=47.892, b=37.012, info="This test case tests a non-square image (4x6).  Make sure you are handling non-square images correctly.")

test3 = findAverageColor(lab3, 0, 0, 4, 2)
print(f"{info} Your test3.png values start=(0, 0) w/ width=4, height=2: (L={test3['L']}, a={test3['a']}, b={test3['b']})")
verify( test3, L=54.244, a=59.314, b=52.980, info="This test case tests a non-square image (4x2).  Make sure you are handling non-square images correctly.")

test3 = findAverageColor(lab3, 2, 2, 2, 2)
print(f"{info} Your test3.png values start=(2, 2) w/ width=2, height=2: (L={test3['L']}, a={test3['a']}, b={test3['b']})")
verify( test3, L=26.057, a=18.195, b=-4.505, info="This test case tests starting at a position other than (0, 0).  Make sure you are starting at (x, y) instead of (0, 0).")

test3 = findAverageColor(lab3, 2, 2, 2, 1)
print(f"{info} Your test3.png values start=(2, 2) w/ width=2, height=1: (L={test3['L']}, a={test3['a']}, b={test3['b']})")
verify( test3, L=16.661, a=4.488, b=-23.667, info="This test case tests starting at a position other than (0, 0).  Make sure you are starting at (x, y) instead of (0, 0).")

test3 = findAverageColor(lab3, 3, 3, 1, 3)
print(f"{info} Your test3.png values start=(3, 3) w/ width=1, height=3: (L={test3['L']}, a={test3['a']}, b={test3['b']})")
verify( test3, L=16.661, a=4.488, b=-23.667, info="This test case tests starting at (3, 3) and going to (4, 6) and is the most complex yet.  You passed a lot of simpler test cases, so you have the basic algorithm down but are not calculating the correct value for non-(0, 0), non-square locations.")


## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print()
print(f"{tada} All tests passed! {tada}")
```

## Part 2b: Finding the best match

In Part 1, you saved a csv file of all tile images.  For this part, you will find the best tile image given an `L`, `a`, and `b` value and your DataFrame, passed in as `df`, in the same format as you saved in Part 1.

This function must return a new DataFrame with exactly one row that contains the best match out of all of the images in `df` based on the Euclidean distance away from the provided (`L`, `a`, `b`).  *(You should not remove rows from `df` itself, as the same `df` will be passed to you each time; make sure to assign your smallest one row to a new and differently named DataFrame.)*

<!-- grade_id: cell-cb3ff910dd835a51, points: 0, grade: false, solution: true -->
```python
# Returns the filename for the image that is the best match given an L, a, and b value.
def findBestMatch(df, L, a, b):
    ...

    
```

<!-- grade_id: cell-72eafa6774690dd4, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Part 2b ==
# - This read-only cell contains test cases for your previous cell(s).
# - If this cell runs without any errors in the output, you PASSED all test cases!
# - If this cell results errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

real_df = pd.DataFrame([
    {'file': 'test.png', 'L': 47.19722525581813, 'a': 49.03421116311881, 'b': 38.60877549417687},
    {'file': 'test2.png', 'L': 54.24409328969397, 'a': 59.3141053878179, 'b': 52.97987993308968},
    {'file': 'test3.png', 'L': 46.41423991872082, 'a': 47.89200069370779, 'b': 37.011986112075455}
])

bestMatch = findBestMatch(real_df, 0, 0, 0)
assert(type(bestMatch) == type(pd.DataFrame())), "findBestMatch must return a DataFrame"
assert(len(bestMatch) == 1), "findBestMatch must return exactly one best match"
assert(bestMatch['file'].values[0] == 'test3.png'), "findBestMatch did not return the best match for test (L=0, a=0, b=0)"

bestMatch = findBestMatch(real_df, 47, 49, 38)
assert(bestMatch['file'].values[0] == 'test.png'), "findBestMatch did not return the best match for test (L=47, a=49, b=38)"

bestMatch = findBestMatch(real_df, 54, 49, 38)
assert(bestMatch['file'].values[0] == 'test.png'), "findBestMatch did not return the best match for test (L=54, a=49, b=38)"

bestMatch = findBestMatch(real_df, 54, 49, 52)
assert(bestMatch['file'].values[0] == 'test2.png'), "findBestMatch did not return the best match for test (L=54, a=49, b=52)"

bestMatch = findBestMatch(real_df, -100, -100, -100)
assert(bestMatch['file'].values[0] == 'test3.png'), "findBestMatch did not return the best match for test (L=-100, a=-100, b=-100)"


## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

## Part 2c: Creating your mosaic!

Finally, the following code uses your image from Part 1, your DataFrame of average colors for each image, the `findAverageColor` function from Part 2a, and `findBestMatch` function from Part 2b to draw a mosaic image!
- Make sure `project1-settings.ipynb` has the correct file name for your base image!
- This code will take a while to run; you may want to change some settings in `project1-settings.ipynb` while testing!

```python
import skimage
import skimage.io

import PIL
from PIL import Image

from os import listdir
from os.path import isfile, join
import sys
import math

# Run the settings notebook
%run project1-settings.ipynb

# Load the saved image data (from Part 1)
print("Loading in saved average image values...")
df = pd.read_csv('tile-images.csv')

# Load the mosaic image:
print(f"Loading the base image `{baseImageFile}`...")
baseImage = Image.open(baseImageFile)
w, h = baseImage.size
lab = imageToLabArray(baseImage, resize=False)

# Calculate the final width/height of the mosaic image based on the tiles that will be included.
# ...this will cut off the very edge of the image to ensure only full tiles are used.
baseImagePixelsPerTile = math.ceil(w / maxTilesAcross)
w_tiles = int(w / baseImagePixelsPerTile)
h_tiles = int(h / baseImagePixelsPerTile)
w = w_tiles * baseImagePixelsPerTile
h = h_tiles * baseImagePixelsPerTile

# Create a blnak image of the correct size to draw the final mosaic on:
baseImage = Image.new('RGB', (outputSize * w_tiles, outputSize * h_tiles))

# Store images used to speed up processing (often once an image is used once it will be used again):
cache = {}

print(f"Creating your moasic ({w_tiles} x {h_tiles} = {w_tiles * h_tiles} total tiles)...")
for x in range(0, w, baseImagePixelsPerTile):
    for y in range(0, h, baseImagePixelsPerTile):
        # Find the average color for the current tile:
        tileAvgColor = findAverageColor(lab, x, y, baseImagePixelsPerTile, baseImagePixelsPerTile)
        
        # Find the best file match:
        df_bestMatch = findBestMatch(df, tileAvgColor['L'], tileAvgColor['a'], tileAvgColor['b'])
        bestFileName = df_bestMatch['file'].values[0]
        
        # load the iamge in and resize it to be a `outputSize` x `outputSize` (or get it from the cache)
        if bestFileName in cache:
            smallTile = cache[bestFileName]
        else:
            tileImage = Image.open(bestFileName)
            smallTile = imageResizeToOutputSize(tileImage)
            cache[bestFileName] = smallTile
        
        # Draw the tile:
        baseImage.paste(smallTile, ( int((x / baseImagePixelsPerTile) * outputSize), int((y / baseImagePixelsPerTile) * outputSize)))
        
    # Print out a progress message:
    completed = int((x / baseImagePixelsPerTile) + 1) * baseImagePixelsPerTile
    pct = completed / (w_tiles * baseImagePixelsPerTile) * 100
    sys.stdout.write(f'\r  ...progress: {completed} / {w_tiles * baseImagePixelsPerTile} ({pct:.2f}%)')

# Save the image as mosaic-hq.jpg
baseImage.save('mosaic-hd.jpg')

# Save a smaller image as mosaic.jpg
d = max(w, h)
factor = d / 4000
if factor <= 1:
    factor = 1

small_w = w / factor
small_h = h / factor    
baseImage = baseImage.resize( (int(small_w), int(small_h)), resample=PIL.Image.LANCZOS )
baseImage.save('mosaic.jpg')


# Print a message:
tada = "\N{PARTY POPPER}"
print("")
print("")
print(f"{tada} MOSAIC COMPLETE! {tada}")
print("- See mosaic-hq.jpg to see your HQ moasic or just mosaic.jpg for a smaller image of the same file.")
```

## Have Fun!

Use Part 1 whenever you change your tile images and Part 2 to build image mosaics -- you will forever be able to build image mosaics now! :D