# Welcome to the University of Illinois administrative office!

We are so thankful that you agreed to take a look at our salary data. We have a dataset for you that we need you to perform some analysis on and create some plots that we can show the state to get funding for another year. Further details follow below.

Thanks again!
 
\- Chancellor Kelly

## Lab: Plots

In this lab, you will:

- Work with real UIUC salary data to explore some of the statistics that we talked about in class: mean, median, mode, standard deviation, and quartiles!
- Practice creating plots to visualize quantitative data: boxplots, scatter plots, and histograms

# 0. Import the Data

In the same directory as this notebook find the dataset that is provided.  You know what to do to import the libraries and read the data! :)

<!-- grade_id: cell-27971beecac131cc, points: 0, grade: false, solution: true -->
```python

```

# 1. The Basics


## Exploring measures of center: Mean, Median, and Standard Deviation

The first thing we need from you is to get us some basics statistics at both a university and department level. We need to know:
1. the number of employees at UIUC
2. the mean salary
3. the median salary
4. the standard deviation

<!-- grade_id: cell-419796a0b415e588, points: 0, grade: false, solution: true -->
```python
university_num_emp = ...
university_num_emp


```

<!-- grade_id: cell-4dbbeef1022a20e4, points: 0, grade: false, solution: true -->
```python
university_mean_sal = ...
university_mean_sal


```

<!-- grade_id: cell-727195c3fef2cf78, points: 0, grade: false, solution: true -->
```python
university_median_sal = ...
university_median_sal


```

<!-- grade_id: cell-35e651d14a628e50, points: 0, grade: false, solution: true -->
```python
university_std_sal = ...
university_std_sal


```

<!-- grade_id: cell-04e624817d681f8e, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1.1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs with a message (with the emoji showing), you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.
x = university_num_emp + university_mean_sal + university_median_sal + university_std_sal
y = ((university_mean_sal - university_median_sal) * university_std_sal) / university_num_emp

assert( abs(x - 246964.21165107167) < 1 ), "It looks like your calculations are off?"
assert( abs(y - 221601.415478511) < 1 ), "It looks like your calculations are off?"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

## Department-level Data

Now we need the same data at a department level. In other words, you want each of those statistics for each department.  There was a certain function you learned a bit ago that will help you **group** all of the people in each department so you can find departmental aggregates.

Instead of just grouping by a single aggregate, use **Python list syntax** to aggregate to find the `count`, `mean`, and `median`.  After that, select only the `'Primary Department'` and `'Salary'` data.

*(You can view technical docs here:
https://pandas.pydata.org/pandas-docs/stable/reference/groupby.html)*

<!-- grade_id: cell-72189b308296916b, points: 0, grade: false, solution: true -->
```python
group = ...
agg_df = ...



# Display 5 random rows to make sure it works:
agg_df.sample(5)
```

<!-- grade_id: cell-d4a69c902629304a, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1.2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs with a message (with the emoji showing), you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.
assert( len(agg_df) == 140 ), "You should have 140 different departments.  Check to make sure you grouped the data?"
assert( "Salary" in agg_df ), "You must have a salary output"
assert( "mean" in agg_df["Salary"] ), "Make sure you have a mean in your salary"
assert( "count" in agg_df["Salary"] ), "Make sure you have a count in your salary"
assert( "median" in agg_df["Salary"] ), "Make sure you have a median in your salary"
assert( len(agg_df.columns) == 4 ), "Make sure you have ONLY the salary data and the department"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

Your boss is thrilled with the numbers you've come up with! Now she's ready to go to the state to ask for money. To know how much to ask for, she wants you to **add a column** to the department-grouped DataFrame that is equal to the **mean** salary of that department multiplied by the number of people in that department. She wants you to provide her with this table and the sum of all of the values in this column.

Hop to it!

<!-- grade_id: cell-4532891e3738e59e, points: 0, grade: false, solution: true -->
```python
agg_df['MeanMulCount'] = ...



# Display 5 random rows to make sure it works:
agg_df.sample(5)
```

<!-- grade_id: cell-6da6f358c0acaf48, points: 0, grade: false, solution: true -->
```python
# calculate the total sum of all of the values in the column you added
total_money_needed_mean = ...


```

<!-- grade_id: cell-f199dd0452902f41, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1.3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs with a message (with the emoji showing), you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.

assert( abs(total_money_needed_mean - 358852459) < 1 ), "Looks like your multiplication did not work properly!"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

# 2. Estimations

Great! Your boss is ready to go to the state with this number. Is this a good estimate of the amount of money she will need?

Let's calculate the standard deviation for each department to find out.  Remember, SD measures the spread around the average.  The larger the SD, the more spread out the numbers are from the average.

Add another column to your dataframe with the departmental data which holds the value of the standard deviation of the salaries for that department.  *(Hint: This may require you to re-group your original `df`)*

<!-- grade_id: cell-dca5763613817466, points: 0, grade: false, solution: true -->
```python
# add that column here!
agg_df = ...



# Display 5 random rows to make sure it works:
agg_df.sample(5)
```

Are there any departments for which your boss' way of approximating the total salary would be fairly accurate? Are there deparments where it would not be an okay approximation?  *Hint*: It might be helpful to find the largest standard deviation!

<!-- grade_id: cell-8d360389fd08fc8f, points: 0, grade: false, solution: true -->
```python

```

## EDIT THIS CELL:

- One department where it would be an okay approximation is:

- One department where it would not be an okay approximation is:

Well, I believe you that the mean salary isn't always a good approximation, but your boss doesn't. I think she is more of a visual learner. To help prove our point, let's make a histogram of the salaries in the computer science department.

## 2.1 Histograms
Create a histogram of the salaries within the computer science department. Pandas has a nice function that helps you accomplish this: https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.hist.html

Hint: in order to apply this function, you will need a DataFrame `cs_salary_df` with just the data for Computer Science professors.

Play around with putting different values as the *bin* parameter to the function. Which value would you select to present the data to your boss?  Remember, there is no right or wrong number of bins.  You want to pick an amount that allows you to see the overall shape of the histogram.

<!-- grade_id: cell-12e1d9777e55620b, points: 0, grade: false, solution: true -->
```python
cs_salary_df = ...

# ...and then create your histogram:



```

<!-- grade_id: cell-5ad401ab98278a69, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2.1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs with a message (with the emoji showing), you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.
assert(cs_salary_df["Salary"].sum() == 13323600), "You are not showing the correct data in the histogram"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

## 2.2 Boxplots

Ok, you've convinced her. Your boss no longer thinks that the average is a good approximation for each department. To explain to the state why this isn't the case, she wants you to create a plot that shows the distribution of the salaries for each department that is relatively compact (i.e., not 100 tiny histograms). A collection of box plots would work quite nicely!  This allows us to see different quartiles and if the departments have outliers!

But first, let's just make one box plot. Again, pandas has a nifty function that can help you do this: https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.boxplot.html

Make a box plot for the computer science department salaries.

<!-- grade_id: cell-bc5e651b041152a9, points: 0, grade: false, solution: true -->
```python
# Create your box plot of the CS salaries:



```

## 2.3 New Estimating Method

Your boss notices some outliers in your boxplot and wants to try another estimation strategy to account for them. She has heard that the median is less likely to be distorted by extreme data points and wants to try using that for salary estimation instead of the mean.

Just like with the mean estimator, she wants you to **add a column** to the department-grouped DataFrame that is equal to the **median** salary of that department multiplied by the number of people in that department. She wants you to provide her again with this table and the sum of all of the values in this column.

<!-- grade_id: cell-ce6f91de7935befc, points: 0, grade: false, solution: true -->
```python
agg_df['MedianMulCount'] = ...



# Display 5 random rows to make sure it works:
agg_df.sample(5)
```

<!-- grade_id: cell-5be4d8ec77de1777, points: 0, grade: false, solution: true -->
```python
# calculate the total sum of all of the values in the column you added
total_money_needed_med = ...


```

```python
## == TEST CASES for Puzzle 2.3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs with a message (with the emoji showing), you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.

assert( abs(total_money_needed_med - 332199238) < 1 ), "Looks like your multiplication did not work properly!"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

# 3. Confronting your boss
After spending days creating plots, you finally decide to tell your boss that there is an easier way to find out how much money you should request from the state to pay university salaries. Find the exact total amount of money that will be needed to pay everyone in the dataset.

<!-- grade_id: cell-674a9d3767b37b18, points: 0, grade: false, solution: true -->
```python
# your code goes here!
total_money = ...


```

<!-- grade_id: cell-a52dae832d7423fb, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs with a message (with the emoji showing), you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.

assert(abs(total_money - 358852459) < 1), "Are you sure the total is correct?"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

# 4. Estimation Error
Your boss is very pleased with your data science knowledge! She would like to see by how much your original estimates were off. Find the exact error (difference) of the mean estimation method and the exact amount needed and then repeat for the median estimate method.

<!-- grade_id: cell-69ce1c4a9f719004, points: 0, grade: false, solution: true -->
```python
#calculate the error of the mean estimation method
mean_estimate_error = ...


```

<!-- grade_id: cell-e593b6cc6e5b2cb9, points: 0, grade: false, solution: true -->
```python
#calculate the error of the mean estimation method
med_estimate_error = ...


```

```python
## == TEST CASES for Puzzle 4 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs with a message (with the emoji showing), you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.

assert(abs(mean_estimate_error - (-5.9605e-08)) < 1.0e-07), "Are you sure the error for the mean estimate is correct?"
assert(abs(med_estimate_error - (-26653221)) < 1), "Are you sure the error for the median estimate is correct?"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

## EDIT THIS CELL:

- The method that resulted with the least error was:
- With an error of:

## 4. Your Department

You're awesome!  Now find something interesting about your home department... maybe a box plot or a scatter plot, do whatever interests you, it does not have to be complex (but totally can be!)

<!-- grade_id: cell-539333ff16c3dfff, points: 1, grade: true, solution: true -->
```python

```

## Submit Your Work

You're almost done -- congratulations! 🎉

You need to do a few more things:

1. Save your work.  To do this, create a **notebook checkpoint** by using the menu within the notebook to go **File -> Save and Checkpoint**


2. After you have saved and checkpointed, exit this notebook by going to **File -> Close and Halt**


3. Once you have closed your notebook select **Quit** on the **Jupyter main page** to stop the Jupyter server


4. Return to the Data Science Discovery page on how to use git to turn this notebook into the course!