# Lab Lists and Dictionaries

Hello and welcome to the CIA. Henceforth, you will be addressed by your codename: "Agent BlackCobra". The information about your assignment is highly classified, and should not be shared with anyone other than your supervisors under any circumstances.

We have intercepted a note from the government of Kerplakistan intended for its military located on a submarine in the Pacific Ocean. The problem is, the message is encoded using a caesar cypher (https://en.wikipedia.org/wiki/Caesar_cipher) We know that the message is dealing with the raid of the SpaceX headquarters, as the Kerplakistinian president recently commented "that Elon Musk dude is crazy, SpaceX has gotta go", but we don't know when the attack is happening.

Please use your programming skills to decode the message, Agent BlackCobra; the fate of Elon Musk rests in your hands.

- Complete all puzzles and submit this lab before Sunday at 11:59pm

## Concept Review: Lists in Python

Several different times during Data Science DISCOVERY, you have seen the use of **lists** in Python.  A list allow us to combine multiple elements into a single variable.

For example, the following code stores a list of the first ten prime numbers:
```
primes = [2, 3, 5, 7, 11, 13, 17, 19, 23, 29]
```

And the following list stores five Game of Thrones characters:
```
got_chars = ['Daenerys Targaryen', 'Jon Snow', 'Gregor Clegane', 'Cersei Lannister', 'Arya Stark']
```

Elements in a list can be indexed just like a DataFrame, except that the first element is always 0, the second element is always 1, etc.

- `primes[0]` returns the first element ("index 0") of `primes` -- the number `2`
- `primes[2]` returns the third element ("index 2") of `primes` -- the number `5`
- `got_chars[1]` returns the second element ("index 1") of `got_chars` -- `'Jon Snow'`

Try out a few yourself:

```python
primes = [2, 3, 5, 7, 11, 13, 17, 19, 23, 29]
got_chars = ['Daenerys Targaryen', 'Jon Snow', 'Gregor Clegane', 'Cersei Lannister', 'Arya Stark']
```

## Concept: Dictionaries

The second concept core to this lab is that of a dictionary.  You have already used dictionaries extensively with DataFrames, but we are going to focus on a simple Python dictionary.

A dictionary is an association between a key and a value -- *a dictionary is just a lookup table where we store information to get back later.*  All values stored are stored by a **key**.  For example, your NetID is a lookup value (or **key**) to your name.

- `'waf'` could look up `'Wade Fagen-Ulmschneider'`
- `'kflan'` could look up `'Karle Flanagan'`

To do this in Python, a dictionary in Python takes three steps:

1. Create a empty dictionary with `{}`.  For example:

```
dict = {}
```

2. Add values to the dictionary using the same syntax as a list, except our index is now the **key**:

```
dict['waf'] = 'Wade Fagen-Ulmschneider'
dict['kflan'] = 'Karle Flanagan'
```

3. You can retrieve the values by reading from the dictionary:

```
dict['waf']   # Returns 'Wade Fagen-Ulmschneider'
```

Try out a few yourself:

<!-- grade_id: cell-c1d17ab92f582e01, points: 0, grade: false, solution: true -->
```python
dict = {}

dict['waf'] = 'Wade Fagen-Ulmschneider'
dict['kflan'] = 'Karle Flanagan'
# TODO: Add yourself to the dictionary!

dict


```

<!-- grade_id: cell-8af9642a5ca090fc, points: 1, grade: true, solution: false -->
```python
## == TEST CASES ==
# - This read-only cell contains test cases for your previous cell(s).
assert(len(dict) >= 3), "You must add a new entry to the dictionary."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()
```

## Overall Strategy

A ceasar cypher is when you take all of the letters of a message and rotate them by a shift amount. If the shift parameter is 2, then all of the C's in the original message will become A's in the cypertext (aka, the encrypted version). See the Wikipedia example for more information: https://en.wikipedia.org/wiki/Caesar_cipher.

We know that the Kerplakistinian note was encoded with a caesar cypher, but we don't know what the shift amount was. How can we figure it out? By trying all 26 possible shift amounts (we will only use UPPERCASE letters)!


## Puzzle 1: helper for decoding messages
First, let's write a function that helps us decode a message, given a certain shift amount.

### Pre-Puzzle 1.1: Shifting `A` by 2

If you are going to shift `A` by 2, what is the result?

- (Solve using indexing and list syntax)

<!-- grade_id: cell-5f849453cf430e62, points: 0, grade: false, solution: true -->
```python
alpha = ["A","B","C","D","E","F","G","H","I","J","K","L","M","N","O","P","Q","R","S","T","U","V","W","X","Y","Z"]
A_shiftBy2 = ...


```

<!-- grade_id: cell-305fdaa3ab0c6f32, points: 1, grade: true, solution: false -->
```python
## == TEST CASES ==
# - This read-only cell contains test cases for your previous cell(s).
assert(A_shiftBy2 == "C"), "Shifting `A` by two should be `C`."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()
```

### Pre-Puzzle 1.2: Shifting `A` by `shift` amount

Instead of shifting just by 2, complete a function that shifts `'A'` by the amount stored in the function parameter `shift`.

*HINT1: Use the list `alpha` to help out!*

*HINT2: What happens when 'shift' parameter is large and greater than length of possible alphabet set? If we shift A by 27, then it should "loop back" and return B*



#### Helpful Concept: Modulo operation
Finds the remainder of a division between two positive integers
```
 1 % 2    # returns 1
12 % 5    # returns 2
14 % 7    # returns 0
27 % 26   # returns 1
```


<!-- grade_id: cell-cfcba9163460c169, points: 0, grade: false, solution: true -->
```python
alpha = ["A","B","C","D","E","F","G","H","I","J","K","L","M","N","O","P","Q","R","S","T","U","V","W","X","Y","Z"]

# Return 'A' shifted by `shift`.
def shiftA(shift):
    ...
    

```

#### Testing `shiftA` - Checking Your Work!

```python
shiftA(1)  # Should be 'B'
```

```python
shiftA(6)  # Should be 'G'
```

```python
shiftA(0)  # Should be 'A'
```

```python
shiftA(0) + shiftA(22) + shiftA(4) + shiftA(18) + shiftA(14) + shiftA(12) + shiftA(4) + '!'
```

<!-- grade_id: cell-bcfa5c84909d7d6c, points: 1, grade: true, solution: false -->
```python
## == TEST CASES ==
# - This read-only cell contains test cases for your previous cell(s).
assert(shiftA(0)  == "A"), f"A shift of 0 must be A -- your algorithm got: {shiftA(0)}"
assert(shiftA(22) == "W"), f"A shift of 22 must be W -- your algorithm got: {shiftA(22)}"
assert(shiftA(4)  == "E"), f"A shift of 4 must be E -- your algorithm got: {shiftA(4)}"
assert(shiftA(18) == "S"), f"A shift of 18 must be S -- your algorithm got: {shiftA(18)}"
assert(shiftA(14) == "O"), f"A shift of 14 must be O -- your algorithm got: {shiftA(14)}"
assert(shiftA(12) == "M"), f"A shift of 12 must be M -- your algorithm got: {shiftA(12)}"
assert(shiftA(4)  == "E"), f"A shift of 4 must be E -- your algorithm got: {shiftA(4)}"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()
```

## Puzzle 2

GOAL: Write a function that takes in one parameter: a shift amount between [0, 25] inclusive. The function should return a dictionary that contains the letters A-Z and what they would be in the original (decoded) message with the given shift amount used during encryption.

Example, if the shift amount is 3, then the function should return a dictionary of the following:
```
{"A": "D", "B": "E", "C": "F" ... "X": "A", "Y": "B", "Z": "C"}
```

<!-- grade_id: cell-e37c76fc14bdc94a, points: 0, grade: false, solution: true -->
```python

```

#### Testing `createDictionary`

Let's make sure your function works before we go further!  Check to make sure you get the expected results:

```python
# Print out the whole dictionary, just to see what's in it:
createDictionary(3)
```

```python
d = createDictionary(3)
d['A']  # Should be 'A' shifted by 3 == 'D'
```

```python
d = createDictionary(3)
d['C']  # Should be 'C' shifted by 3 == 'F'
```

```python
d = createDictionary(3)
d['Z']  # Should be 'Z' shifted by 3 == 'C'
```

```python
d = createDictionary(1)
d['Z']  # Should be 'Z' shifted by 1 == 'A'
```

<!-- grade_id: cell-0b65e7c4819d2b25, points: 1, grade: true, solution: false -->
```python
## == TEST CASES ==
# - This read-only cell contains test cases for your previous cell(s).
d = createDictionary(3)
assert(d['A']  == 'D'), f"A dictionary with a shift of 3 must translate A to D -- you got: {d['A']}"
assert(d['C']  == 'F'), f"A dictionary with a shift of 3 must translate C to F -- you got: {d['C']}"
assert(d['Z']  == 'C'), f"A dictionary with a shift of 3 must translate Z to C -- you got: {d['Z']}"

d = createDictionary(1)
assert(d['Z']  == 'A'), f"A dictionary with a shift of 1 must translate Z to A -- you got: {d['Z']}"


## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()
```

## Puzzle 3: Decoding the message

Using your `createDictionary` function, we have provided a function that, given the parameters of the encoded message and a shift amount, returns what that message decodes to via the given shift amount.  The message will be given as a list of the words of the message and the function returns the decoded list:

```python
# This function is complete -- run it and see it work! :)
def decodeGivenShift(msg, shift):
    decoded = []
    dictionary = createDictionary(shift)
    for word in msg:
        decodedWord = ""
        for letter in word:
            decodedWord += dictionary[letter]
        decoded.append(decodedWord)
    return decoded

decodeGivenShift(["HVWG", "WG", "O", "ASGGOUS"], 12)

```

## Puzzle 3.1: Decoding the RIGHT message

Now that we have a decode function given a shift amount, how can we find out which shift amount is right? Easy, try all of the shift amounts, and see which one gives us a message back in English!

We would like you to programmatically detect which message is in English. (How can you do this? Hint: what or who is the known target for the attack? How can we use this information to check if the message was decoded correctly?)

The actual message you need to decode is given in the next cell.

<!-- grade_id: cell-5ca9785aeb50f488, points: 0, grade: false, solution: true -->
```python
def decode(msg):
    ...



kerplakistanMessage = ["KHB", "ERE", "ZH", "DUH", "JRLQJ", "WR", "KLW", "XS", "ORV", "DQJHOHV", "DQG", "VSDFHA", "WRPRUURZ", "DW", "R", "IRXUWHHQ", "KXQGUHG", "WKDW", "HORQ", "GXGH", "LV", "FUDB", "DQG", "ZH", "KDYH", "JRW", "WR", "GR", "VRPHWKLQJ", "DERXW", "LW", "KDVKWDJ", "DLQW", "QR", "PDQ", "WRXFKLQJ", "WKDW", "EDOO", "RI", "FKHHVH", "KDVKWDJ", "GRZQ", "ZLWK", "HORQ", "KDVKWDJ", "MHW", "SDFNV", "EH", "ZKDFN"]
' '.join(decode(kerplakistanMessage))
```

## Puzzle 3.2: A second note

Agent BlackCobra, our operatives intersected a second note from the Kerplakistinian government. We know what they are doing with Elon, so we have no idea what this one is regarding. Can you still find a way to decode the message?

There are several heuristics you can use to solve this issue. For instance, chances are the message contains at least one of the words "the, and, an, but, so". However, we want you to solve it by visualizing the frequency of the letters in the message!

## A touch of background
It turns out that most English texts have relatively consistently distributed letter usage. Traditional English text usually has a letter distribution that looks something like this: https://en.wikipedia.org/wiki/Letter_frequency#/media/File:English_letter_frequency_(alphabetic).svg

Let's create a program that creates a graph of letter frequencies for us.

## Puzzle 4: Visualizing letter frequencies
In order to make the visualization, we have to count letter frequencies first.

GOAL: Create a function that, given a list of strings like before, returns a dictionary where the keys are letters and the values are number of times that letter appears in the entire message.

Example:
```
msg = ["MY", "CAT", "IS", "COOL"]
```
returns: `{'A': 1, 'C': 2, ... 'Y': 1}`

<!-- grade_id: cell-9f1a8da1982ef2d0, points: 0, grade: false, solution: true -->
```python
def countLetters(msg):
    dictionary = {}
    
    ...
    
    return dictionary
    


message = ["MY", "CAT", "IS", "COOL"]
print(countLetters(message))
```

<!-- grade_id: cell-c7e7567f7610d776, points: 1, grade: true, solution: false -->
```python
## == TEST CASES ==
# - This read-only cell contains test cases for your previous cell(s).
message = ["MY", "CAT", "IS", "COOL"]
result = countLetters(message)
assert( result['M'] == 1 ), f"The message \"MY CAT IS COOL\" has one 'M' -- you returned {result['M']}."
assert( result['C'] == 2 ), f"The message \"MY CAT IS COOL\" has two 'C's -- you returned {result['C']}."

message = ["DATA", "SCIENCE", "DISCOVERY"]
result = countLetters(message)
assert( result['Y'] == 1 ), f"The message \"DATA SCIENCE DISCOVERY\" has one 'Y' -- you returned {result['Y']}."
assert( result['E'] == 3 ), f"The message \"DATA SCIENCE DISCOVERY\" has three 'E's -- you returned {result['E']}."
assert( result['C'] == 3 ), f"The message \"DATA SCIENCE DISCOVERY\" has three 'C's -- you returned {result['C']}."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

## Create the visualization

We have provided a function that creates a visualization of the letter frequency of a message that is decoded with a given shift amount. The function takes in two parameters: an encoded message and a shift amount to try, and it will produce a frequency plot of the letters in the decoded message and returning the DataFrame

```python
import pandas as pd
import matplotlib.pyplot as plt

alpha = ["A","B","C","D","E","F","G","H","I","J","K","L","M","N","O","P","Q","R","S","T","U","V","W","X","Y","Z"]
message = ["HVWG", "WG", "O", "ASGGOUS"]

def createVis(msg, shift):
    decodedMsg = decodeGivenShift(msg, shift)
    dictionary = countLetters(decodedMsg)
    data = []
    for key in dictionary:
        data.append({'letter': key, 'count': dictionary[key]} )
    df = pd.DataFrame(data)
    df = df.sort_values(by='letter')
    df.plot.bar(x = 'letter', y = 'count')
    return df
        
df = createVis(message, 12)
```

## What is the right shift value?
Now try manually passing parameters into the createVis function and see which shift value is correct (e.g., which parameter creates a distribution that matches this one https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/English_letter_frequency_%28alphabetic%29.svg/1200px-English_letter_frequency_%28alphabetic%29.svg.png)

Note that the correct shift value will create a distribution that **very clearly** mimics the wikipedia image. It will be obvious which solution is correct. (Hint, if you start with shift value 0 and you try them in order, if you get to 12 you've passed the correct answer)

<!-- grade_id: cell-3edc3f33d68f4388, points: 0, grade: false, solution: true -->
```python
# Using the createViz visualization, find the correct shift value:
shiftValue = 0



# This will send the message through `createViz`:
newMessage = ['ZSNW','QGM','JWSV','SDAUW','AF','OGFVWJDSFV','KMUZ','S','YJWSL','TGGC','EQ','XSNGJALW','IMGLW','AK','SDAUW','OSK','TWYAFFAFY','LG','YWL','NWJQ','LAJWV','GX','KALLAFY','TQ','ZWJ','KAKLWJ','GF','LZW','TSFC','SFV','GX','ZSNAFY','FGLZAFY','LG','VG','GFUW','GJ','LOAUW','KZW','ZSV','HWWHWV','AFLG','LZW','TGGC','ZWJ','KAKLWJ','OSK','JWSVAFY','TML','AL','ZSV','FG','HAULMJWK','GJ','UGFNWJKSLAGFK','AF','AL','SFV','OZSL','AK','LZW','MKW','GX','S','TGGC','LZGMYZL','SDAUW','OALZGML','HAULMJWK','GJ','UGFNWJKSLAGF','KG','KZW','OSK','UGFKAVWJAFY','AF','ZWJ','GOF','EAFV','SK','OWDD','SK','KZW','UGMDV','XGJ','LZW','ZGL','VSQ','ESVW','ZWJ','XWWD','NWJQ','KDWWHQ','SFV','KLMHAV','OZWLZWJ','LZW','HDWSKMJW','GX','ESCAFY','S','VSAKQUZSAF','OGMDV','TW','OGJLZ','LZW','LJGMTDW','GX','YWLLAFY','MH','SFV','HAUCAFY','LZW','VSAKAWK','OZWF','KMVVWFDQ','S','OZALW','JSTTAL','OALZ','HAFC','WQWK','JSF','UDGKW','TQ','ZWJ','LZWJW','OSK','FGLZAFY','KG','NWJQ','JWESJCSTDW','AF','LZSL','FGJ','VAV','SDAUW','LZAFC','AL','KG','NWJQ','EMUZ','GML','GX','LZW','OSQ','LG','ZWSJ','LZW','JSTTAL','KSQ','LG','ALKWDX','GZ','VWSJ','GZ','VWSJ','A','KZSDD','TW','DSLW','OZWF','KZW','LZGMYZL','AL','GNWJ','SXLWJOSJVK','AL','GUUMJJWV','LG','ZWJ','LZSL','KZW','GMYZL','LG','ZSNW','OGFVWJWV','SL','LZAK','TML','SL','LZW','LAEW','AL','SDD','KWWEWV','IMALW','FSLMJSD','TML','OZWF','LZW','JSTTAL','SULMSDDQ','LGGC','S','OSLUZ','GML','GX','ALK','OSAKLUGSL','HGUCWL','SFV','DGGCWV','SL','AL','SFV','LZWF','ZMJJAWV','GF','SDAUW','KLSJLWV','LG','ZWJ','XWWL','XGJ','AL','XDSKZWV','SUJGKK','ZWJ','EAFV','LZSL','KZW','ZSV','FWNWJ','TWXGJW','KWWF','S','JSTTAL','OALZ','WALZWJ','S','OSAKLUGSLHGUCWL','GJ','S','OSLUZ','LG','LSCW','GML','GX','AL','SFV','TMJFAFY','OALZ','UMJAGKALQ','KZW','JSF','SUJGKK','LZW','XAWDV','SXLWJ','AL','SFV','XGJLMFSLWDQ','OSK','BMKL','AF','LAEW','LG','KWW','AL','HGH','VGOF','S','DSJYW','JSTTALZGDW','MFVWJ','LZW','ZWVYW','AF','SFGLZWJ','EGEWFL','VGOF','OWFL','SDAUW','SXLWJ','AL','FWNWJ','GFUW','UGFKAVWJAFY','ZGO','AF','LZW','OGJDV','KZW','OSK','LG','YWL','GML','SYSAF','LZW','JSTTALZGDW','OWFL','KLJSAYZL','GF','DACW','S','LMFFWD','XGJ','KGEW','OSQ','SFV','LZWF','VAHHWV','KMVVWFDQ','VGOF','KG','KMVVWFDQ','LZSL','SDAUW','ZSV','FGL','S','EGEWFL','LG','LZAFC','STGML','KLGHHAFY','ZWJKWDX','TWXGJW','KZW','XGMFV','ZWJKWDX','XSDDAFY','VGOF','S','NWJQ','VWWH','OWDD','WALZWJ','LZW','OWDD','OSK','NWJQ','VWWH','GJ','KZW','XWDD','NWJQ','KDGODQ','XGJ','KZW','ZSV','HDWFLQ','GX','LAEW','SK','KZW','OWFL','VGOF','LG','DGGC','STGML','ZWJ','SFV','LG','OGFVWJ','OZSL','OSK','YGAFY','LG','ZSHHWF','FWPL','XAJKL','KZW','LJAWV','LG','DGGC','VGOF','SFV','ESCW','GML','OZSL','KZW','OSK','UGEAFY','LG','TML','AL','OSK','LGG','VSJC','LG','KWW','SFQLZAFY','LZWF','KZW','DGGCWV','SL','LZW','KAVWK','GX','LZW','OWDD','SFV','FGLAUWV','LZSL','LZWQ','OWJW','XADDWV','OALZ','UMHTGSJVK','SFV','TGGCKZWDNWK','ZWJW','SFV','LZWJW','KZW','KSO','ESHK','SFV','HAULMJWK','ZMFY','MHGF','HWYK','KZW','LGGC','VGOF','S','BSJ','XJGE','GFW','GX','LZW','KZWDNWK','SK','KZW','HSKKWV','AL','OSK','DSTWDDWV','GJSFYW','ESJESDSVW','TML','LG','ZWJ','YJWSL','VAKSHHGAFLEWFL','AL','OSK','WEHLQ','KZW','VAV','FGL','DACW','LG','VJGH','LZW','BSJ','XGJ','XWSJ','GX','CADDAFY','KGEWTGVQ','KG','ESFSYWV','LG','HML','AL','AFLG','GFW','GX','LZW','UMHTGSJVK','SK','KZW','XWDD','HSKL','AL','OWDD','LZGMYZL','SDAUW','LG','ZWJKWDX','SXLWJ','KMUZ','S','XSDD','SK','LZAK','A','KZSDD','LZAFC','FGLZAFY','GX','LMETDAFY','VGOF','KLSAJK','ZGO','TJSNW','LZWQDD','SDD','LZAFC','EW','SL','ZGEW','OZQ','A','OGMDVFL','KSQ','SFQLZAFY','STGML','AL','WNWF','AX','A','XWDD','GXX','LZW','LGH','GX','LZW','ZGMKW','OZAUZ','OSK','NWJQ','DACWDQ','LJMW','VGOF','VGOF','VGOF','OGMDV','LZW','XSDD','FWNWJ','UGEW','LG','SF','WFV','A','OGFVWJ','ZGO','ESFQ','EADWK','ANW','XSDDWF','TQ','LZAK','LAEW','KZW','KSAV','SDGMV','A','EMKL','TW','YWLLAFY','KGEWOZWJW','FWSJ','LZW','UWFLJW','GX','LZW','WSJLZ','DWL','EW','KWW','LZSL','OGMDV','TW','XGMJ','LZGMKSFV','EADWK','VGOF','A','LZAFC','XGJ','QGM','KWW','SDAUW','ZSV','DWSJFL','KWNWJSD','LZAFYK','GX','LZAK','KGJL','AF','ZWJ','DWKKGFK','AF','LZW','KUZGGDJGGE','SFV','LZGMYZ','LZAK','OSK','FGL','S','NWJQ','YGGV','GHHGJLMFALQ','XGJ','KZGOAFY','GXX','ZWJ','CFGODWVYW','SK','LZWJW','OSK','FG','GFW','LG','DAKLWF','LG','ZWJ','KLADD','AL','OSK','YGGV','HJSULAUW','LG','KSQ','AL','GNWJ','QWK','LZSLK','STGML','LZW','JAYZL','VAKLSFUWTML','LZWF','A','OGFVWJ','OZSL','DSLALMVW','GJ','DGFYALMVW','ANW','YGL','LG','SDAUW','ZSV','FG','AVWS','OZSL','DSLALMVW','OSK','GJ','DGFYALMVW','WALZWJ','TML','LZGMYZL','LZWQ','OWJW','FAUW','YJSFV','OGJVK','LG','KSQ','HJWKWFLDQ','KZW','TWYSF','SYSAF','A','OGFVWJ','AX','A','KZSDD','XSDD','JAYZL','LZJGMYZ','LZW','WSJLZ','ZGO','XMFFQ','ALDD','KWWE','LG','UGEW','GML','SEGFY','LZW','HWGHDW','LZSL','OSDC','OALZ','LZWAJ','ZWSVK','VGOFOSJV','LZW','SFLAHSLZAWK','A','LZAFC','KZW','OSK','JSLZWJ','YDSV','LZWJW','OSK','FG','GFW','DAKLWFAFY','LZAK','LAEW','SK','AL','VAVFL','KGMFV','SL','SDD','LZW','JAYZL','OGJV','TML','A','KZSDD','ZSNW','LG','SKC','LZWE','OZSL','LZW','FSEW','GX','LZW','UGMFLJQ','AK','QGM','CFGO','HDWSKW','ESSE','AK','LZAK','FWO','RWSDSFV','GJ','SMKLJSDAS','SFV','KZW','LJAWV','LG','UMJLKWQ','SK','KZW','KHGCWXSFUQ','UMJLKWQAFY','SK','QGMJW','XSDDAFY','LZJGMYZ','LZW','SAJ','VG','QGM','LZAFC','QGM','UGMDV','ESFSYW','AL','SFV','OZSL','SF','AYFGJSFL','DALLDW','YAJD','KZWDD','LZAFC','EW','XGJ','SKCAFY','FG','ALDD','FWNWJ','VG','LG','SKC','HWJZSHK','A','KZSDD','KWW','AL','OJALLWF','MH','KGEWOZWJW','VGOF','VGOF','VGOF','LZWJW','OSK','FGLZAFY','WDKW','LG','VG','KG','SDAUW','KGGF','TWYSF','LSDCAFY','SYSAF','VAFSZDD','EAKK','EW','NWJQ','EMUZ','LGFAYZL','A','KZGMDV','LZAFC','VAFSZ','OSK','LZW','USL','A','ZGHW','LZWQDD','JWEWETWJ','ZWJ','KSMUWJ','GX','EADC','SL','LWSLAEW','VAFSZ','EQ','VWSJ','A','OAKZ','QGM','OWJW','VGOF','ZWJW','OALZ','EW','LZWJW','SJW','FG','EAUW','AF','LZW','SAJ','AE','SXJSAV','TML','QGM','EAYZL','USLUZ','S','TSL','SFV','LZSLK','NWJQ','DACW','S','EGMKW','QGM','CFGO','TML','VG','USLK','WSL','TSLK','A','OGFVWJ','SFV','ZWJW','SDAUW','TWYSF','LG','YWL','JSLZWJ','KDWWHQ','SFV','OWFL','GF','KSQAFY','LG','ZWJKWDX','AF','S','VJWSEQ','KGJL','GX','OSQ','VG','USLK','WSL','TSLK','VG','USLK','WSL','TSLK','SFV','KGEWLAEWK','VG','TSLK','WSL','USLK','XGJ','QGM','KWW','SK','KZW','UGMDVFL','SFKOWJ','WALZWJ','IMWKLAGF','AL','VAVFL','EMUZ','ESLLWJ','OZAUZ','OSQ','KZW','HML','AL','KZW','XWDL','LZSL','KZW','OSK','VGRAFY','GXX','SFV','ZSV','BMKL','TWYMF','LG','VJWSE','LZSL','KZW','OSK','OSDCAFY','ZSFV','AF','ZSFV','OALZ','VAFSZ','SFV','KSQAFY','LG','ZWJ','NWJQ','WSJFWKLDQ','FGO','VAFSZ','LWDD','EW','LZW','LJMLZ','VAV','QGM','WNWJ','WSL','S','TSL','OZWF','KMVVWFDQ','LZMEH','LZMEH','VGOF','KZW','USEW','MHGF','S','ZWSH','GX','KLAUCK','SFV','VJQ','DWSNWK','SFV','LZW','XSDD','OSK','GNWJ','SDAUW','OSK','FGL','S','TAL','ZMJL','SFV','KZW','BMEHWV','MH','GF','LG','ZWJ','XWWL','AF','S','EGEWFL','KZW','DGGCWV','MH','TML','AL','OSK','SDD','VSJC','GNWJZWSV','TWXGJW','ZWJ','OSK','SFGLZWJ','DGFY','HSKKSYW','SFV','LZW','OZALW','JSTTAL','OSK','KLADD','AF','KAYZL','ZMJJQAFY','VGOF','AL','LZWJW','OSK','FGL','S','EGEWFL','LG','TW','DGKL','SOSQ','OWFL','SDAUW','DACW','LZW','OAFV','SFV','OSK','BMKL','AF','LAEW','LG','ZWSJ','AL','KSQ','SK','AL','LMJFWV','S','UGJFWJ','GZ','EQ','WSJK','SFV','OZAKCWJK','ZGO','DSLW','ALK','YWLLAFY','KZW','OSK','UDGKW','TWZAFV','AL','OZWF','KZW','LMJFWV','LZW','UGJFWJ','TML','LZW','JSTTAL','OSK','FG','DGFYWJ','LG','TW','KWWF','KZW','XGMFV','ZWJKWDX','AF','S','DGFY','DGO','ZSDD','OZAUZ','OSK','DAL','MH','TQ','S','JGO','GX','DSEHK','ZSFYAFY','XJGE','LZW','JGGX','LZWJW','OWJW','VGGJK','SDD','JGMFV','LZW','ZSDD','TML','LZWQ','OWJW','SDD','DGUCWV','SFV','OZWF','SDAUW','ZSV','TWWF','SDD','LZW','OSQ','VGOF','GFW','KAVW','SFV','MH','LZW','GLZWJ','LJQAFY','WNWJQ','VGGJ','KZW','OSDCWV','KSVDQ','VGOF','LZW','EAVVDW','OGFVWJAFY','ZGO','KZW','OSK','WNWJ','LG','YWL','GML','SYSAF','KMVVWFDQ','KZW','USEW','MHGF','S','DALLDW','LZJWWDWYYWV','LSTDW','SDD','ESVW','GX','KGDAV','YDSKK','LZWJW','OSK','FGLZAFY','GF','AL','WPUWHL','S','LAFQ','YGDVWF','CWQ','SFV','SDAUWK','XAJKL','LZGMYZL','OSK','LZSL','AL','EAYZL','TWDGFY','LG','GFW','GX','LZW','VGGJK','GX','LZW','ZSDD','TML','SDSK','WALZWJ','LZW','DGUCK','OWJW','LGG','DSJYW','GJ','LZW','CWQ','OSK','LGG','KESDD','TML','SL','SFQ','JSLW','AL','OGMDV','FGL','GHWF','SFQ','GX','LZWE','ZGOWNWJ','GF','LZW','KWUGFV','LAEW','JGMFV','KZW','USEW','MHGF','S','DGO','UMJLSAF','KZW','ZSV','FGL','FGLAUWV','TWXGJW','SFV','TWZAFV','AL','OSK','S','DALLDW','VGGJ','STGML','XAXLWWF','AFUZWK','ZAYZ','KZW','LJAWV','LZW','DALLDW','YGDVWF','CWQ','AF','LZW','DGUC','SFV','LG','ZWJ','YJWSL','VWDAYZL','AL','XALLWV','SDAUW','GHWFWV','LZW','VGGJ','SFV','XGMFV','LZSL','AL','DWV','AFLG','S','KESDD','HSKKSYW','FGL','EMUZ','DSJYWJ','LZSF','S','JSLZGDW','KZW','CFWDL','VGOF','SFV','DGGCWV','SDGFY','LZW','HSKKSYW','AFLG','LZW','DGNWDAWKL','YSJVWF','QGM','WNWJ','KSO','ZGO','KZW','DGFYWV','LG','YWL','GML','GX','LZSL','VSJC','ZSDD','SFV','OSFVWJ','STGML','SEGFY','LZGKW','TWVK','GX','TJAYZL','XDGOWJK','SFV','LZGKW','UGGD','XGMFLSAFK','TML','KZW','UGMDV','FGL','WNWF','YWL','ZWJ','ZWSV','LZJGMYZ','LZW','VGGJOSQ','SFV','WNWF','AX','EQ','ZWSV','OGMDV','YG','LZJGMYZ','LZGMYZL','HGGJ','SDAUW','AL','OGMDV','TW','GX','NWJQ','DALLDW','MKW','OALZGML','EQ','KZGMDVWJK','GZ','ZGO','A','OAKZ','A','UGMDV','KZML','MH','DACW','S','LWDWKUGHW','A','LZAFC','A','UGMDV','AX','A','GFDQ','CFGO','ZGO','LG','TWYAF','XGJ','QGM','KWW','KG','ESFQ','GMLGXLZWOSQ','LZAFYK','ZSV','ZSHHWFWV','DSLWDQ','LZSL','SDAUW','ZSV','TWYMF','LG','LZAFC','LZSL','NWJQ','XWO','LZAFYK','AFVWWV','OWJW','JWSDDQ','AEHGKKATDW','LZWJW','KWWEWV','LG','TW','FG','MKW','AF','OSALAFY','TQ','LZW','DALLDW','VGGJ','KG','KZW','OWFL','TSUC','LG','LZW','LSTDW','ZSDX','ZGHAFY','KZW','EAYZL','XAFV','SFGLZWJ','CWQ','GF','AL','GJ','SL','SFQ','JSLW','S','TGGC','GX','JMDWK','XGJ','KZMLLAFY','HWGHDW','MH','DACW','LWDWKUGHWK','LZAK','LAEW','KZW','XGMFV','S','DALLDW','TGLLDW','GF','AL','OZAUZ','UWJLSAFDQ','OSK','FGL','ZWJW','TWXGJW','KSAV','SDAUW','SFV','JGMFV','LZW','FWUC','GX','LZW','TGLLDW','OSK','S','HSHWJ','DSTWD','OALZ','LZW','OGJVK','VJAFC','EW','TWSMLAXMDDQ','HJAFLWV','GF','AL','AF','DSJYW','DWLLWJK','AL','OSK','SDD','NWJQ','OWDD','LG','KSQ','VJAFC','EW','TML','LZW','OAKW','DALLDW','SDAUW','OSK','FGL','YGAFY','LG','VG','LZSL','AF','S','ZMJJQ','FG','ADD','DGGC','XAJKL','KZW','KSAV','SFV','KWW','OZWLZWJ','ALK','ESJCWV','HGAKGF','GJ','FGL','XGJ','KZW','ZSV','JWSV','KWNWJSD','FAUW','DALLDW','ZAKLGJAWK','STGML','UZADVJWF','OZG','ZSV','YGL','TMJFL','SFV','WSLWF','MH','TQ','OADV','TWSKLK','SFV','GLZWJ','MFHDWSKSFL','LZAFYK','SDD','TWUSMKW','LZWQ','OGMDV','FGL','JWEWETWJ','LZW','KAEHDW','JMDWK','LZWAJ','XJAWFVK','ZSV','LSMYZL','LZWE','KMUZ','SK','LZSL','S','JWVZGL','HGCWJ','OADD','TMJF','QGM','AX','QGM','ZGDV','AL','LGG','DGFY','SFV','LZSL','AX','QGM','UML','QGMJ','XAFYWJ','NWJQ','VWWHDQ','OALZ','S','CFAXW','AL','MKMSDDQ','TDWWVK','SFV','KZW','ZSV','FWNWJ','XGJYGLLWF','LZSL','AX','QGM','VJAFC','EMUZ','XJGE','S','TGLLDW','ESJCWV','HGAKGF','AL','AK','SDEGKL','UWJLSAF','LG','VAKSYJWW','OALZ','QGM','KGGFWJ','GJ','DSLWJ','ZGOWNWJ','LZAK','TGLLDW','OSK','FGL','ESJCWV','HGAKGF','KG','SDAUW','NWFLMJWV','LG','LSKLW','AL','SFV','XAFVAFY','AL','NWJQ','FAUW','AL','ZSV','AF','XSUL','S','KGJL','GX','EAPWV','XDSNGMJ','GX','UZWJJQLSJL','UMKLSJV','HAFWSHHDW','JGSKL','LMJCWQ','LGXXWW','SFV','ZGL','TMLLWJWV','LGSKL','KZW','NWJQ','KGGF','XAFAKZWV','AL','GXX','OZSL','S','UMJAGMK','XWWDAFY','KSAV','SDAUW','A','EMKL','TW','KZMLLAFY','MH','DACW','S','LWDWKUGHW','SFV','KG','AL','OSK','AFVWWV','KZW','OSK','FGO','GFDQ','LWF','AFUZWK','ZAYZ','SFV','ZWJ','XSUW','TJAYZLWFWV','MH','SL','LZW','LZGMYZL','LZSL','KZW','OSK','FGO','LZW','JAYZL','KARW','XGJ','YGAFY','LZJGMYZ','LZW','DALLDW','VGGJ','AFLG','LZSL','DGNWDQ','YSJVWF','XAJKL','ZGOWNWJ','KZW','OSALWV','XGJ','S','XWO','EAFMLWK','LG','KWW','AX','KZW','OSK','YGAFY','LG','KZJAFC','SFQ','XMJLZWJ','KZW','XWDL','S','DALLDW','FWJNGMK','STGML','LZAK','XGJ','AL','EAYZL','WFV','QGM','CFGO','KSAV','SDAUW','LG','ZWJKWDX','AF','EQ','YGAFY','GML','SDLGYWLZWJ','DACW','S','USFVDW','A','OGFVWJ','OZSL','A','KZGMDV','TW','DACW','LZWF','SFV','KZW','LJAWV','LG','XSFUQ','OZSL','LZW','XDSEW','GX','S','USFVDW','AK','DACW','SXLWJ','LZW','USFVDW','AK','TDGOF','GML','XGJ','KZW','UGMDV','FGL','JWEWETWJ','WNWJ','ZSNAFY','KWWF','KMUZ','S','LZAFY','SXLWJ','S','OZADW','XAFVAFY','LZSL','FGLZAFY','EGJW','ZSHHWFWV','KZW','VWUAVWV','GF','YGAFY','AFLG','LZW','YSJVWF','SL','GFUW','TML','SDSK','XGJ','HGGJ','SDAUW','OZWF','KZW','YGL','LG','LZW','VGGJ','KZW','XGMFV','KZW','ZSV','XGJYGLLWF','LZW','DALLDW','YGDVWF','CWQ','SFV','OZWF','KZW','OWFL','TSUC','LG','LZW','LSTDW','XGJ','AL','KZW','XGMFV','KZW','UGMDV','FGL','HGKKATDQ','JWSUZ','AL','KZW','UGMDV','KWW','AL','IMALW','HDSAFDQ','LZJGMYZ','LZW','YDSKK','SFV','KZW','LJAWV','ZWJ','TWKL','LG','UDAET','MH','GFW','GX','LZW','DWYK','GX','LZW','LSTDW','TML','AL','OSK','LGG','KDAHHWJQ','SFV','OZWF','KZW','ZSV','LAJWV','ZWJKWDX','GML','OALZ','LJQAFY','LZW','HGGJ','DALLDW','LZAFY','KSL','VGOF','SFV','UJAWV','UGEW','LZWJWK','FG','MKW','AF','UJQAFY','DACW','LZSL','KSAV','SDAUW','LG','ZWJKWDX','JSLZWJ','KZSJHDQ','A','SVNAKW','QGM','LG','DWSNW','GXX','LZAK','EAFMLW','KZW','YWFWJSDDQ','YSNW','ZWJKWDX','NWJQ','YGGV','SVNAUW','LZGMYZ','KZW','NWJQ','KWDVGE','XGDDGOWV','AL','SFV','KGEWLAEWK','KZW','KUGDVWV','ZWJKWDX','KG','KWNWJWDQ','SK','LG','TJAFY','LWSJK','AFLG','ZWJ','WQWK','SFV','GFUW','KZW','JWEWETWJWV','LJQAFY','LG','TGP','ZWJ','GOF','WSJK','XGJ','ZSNAFY','UZWSLWV','ZWJKWDX','AF','S','YSEW','GX','UJGIMWL','KZW','OSK','HDSQAFY','SYSAFKL','ZWJKWDX','XGJ','LZAK','UMJAGMK','UZADV','OSK','NWJQ','XGFV','GX','HJWLWFVAFY','LG','TW','LOG','HWGHDW','TML','ALK','FG','MKW','FGO','LZGMYZL','HGGJ','SDAUW','LG','HJWLWFV','LG','TW','LOG','HWGHDW','OZQ','LZWJWK','ZSJVDQ','WFGMYZ','GX','EW','DWXL','LG','ESCW','GFW','JWKHWULSTDW','HWJKGF','KGGF','ZWJ','WQW','XWDD','GF','S','DALLDW','YDSKK','TGP','LZSL','OSK','DQAFY','MFVWJ','LZW','LSTDW','KZW','GHWFWV','AL','SFV','XGMFV','AF','AL','S','NWJQ','KESDD','USCW','GF','OZAUZ','LZW','OGJVK','WSL','EW','OWJW','TWSMLAXMDDQ','ESJCWV','AF','UMJJSFLK','OWDD','ADD','WSL','AL','KSAV','SDAUW','SFV','AX','AL','ESCWK','EW','YJGO','DSJYWJ','A','USF','JWSUZ','LZW','CWQ','SFV','AX','AL','ESCWK','EW','YJGO','KESDDWJ','A','USF','UJWWH','MFVWJ','LZW','VGGJ','KG','WALZWJ','OSQ','ADD','YWL','AFLG','LZW','YSJVWF','SFV','A','VGFL','USJW','OZAUZ','ZSHHWFK','KZW','SLW','S','DALLDW','TAL','SFV','KSAV','SFPAGMKDQ','LG','ZWJKWDX','OZAUZ','OSQ','OZAUZ','OSQ','ZGDVAFY','ZWJ','ZSFV','GF','LZW','LGH','GX','ZWJ','ZWSV','LG','XWWD','OZAUZ','OSQ','AL','OSK','YJGOAFY','SFV','KZW','OSK','IMALW','KMJHJAKWV','LG','XAFV','LZSL','KZW','JWESAFWV','LZW','KSEW','KARW','LG','TW','KMJW','LZAK','YWFWJSDDQ','ZSHHWFK','OZWF','GFW','WSLK','USCW','TML','SDAUW','ZSV','YGL','KG','EMUZ','AFLG','LZW','OSQ','GX','WPHWULAFY','FGLZAFY','TML','GMLGXLZWOSQ','LZAFYK','LG','ZSHHWF','LZSL','AL','KWWEWV','IMALW','VMDD','SFV','KLMHAV','XGJ','DAXW','LG','YG','GF','AF','LZW','UGEEGF','OSQ','KG','KZW','KWL','LG','OGJC','SFV','NWJQ','KGGF','XAFAKZWV','GXX','LZW','USCW']
df = createVis(newMessage, shiftValue)
```

<!-- grade_id: cell-da17050954630c0e, points: 0, grade: true, solution: false -->
```python
## == TEST CASES ==
# - This read-only cell contains test cases for your previous cell(s).
newMessage = ['ZSNW','QGM','JWSV','SDAUW','AF','OGFVWJDSFV','KMUZ','S','YJWSL','TGGC','EQ','XSNGJALW','IMGLW','AK','SDAUW','OSK','TWYAFFAFY','LG','YWL','NWJQ','LAJWV','GX','KALLAFY','TQ','ZWJ','KAKLWJ','GF','LZW','TSFC','SFV','GX','ZSNAFY','FGLZAFY','LG','VG','GFUW','GJ','LOAUW','KZW','ZSV','HWWHWV','AFLG','LZW','TGGC','ZWJ','KAKLWJ','OSK','JWSVAFY','TML','AL','ZSV','FG','HAULMJWK','GJ','UGFNWJKSLAGFK','AF','AL','SFV','OZSL','AK','LZW','MKW','GX','S','TGGC','LZGMYZL','SDAUW','OALZGML','HAULMJWK','GJ','UGFNWJKSLAGF','KG','KZW','OSK','UGFKAVWJAFY','AF','ZWJ','GOF','EAFV','SK','OWDD','SK','KZW','UGMDV','XGJ','LZW','ZGL','VSQ','ESVW','ZWJ','XWWD','NWJQ','KDWWHQ','SFV','KLMHAV','OZWLZWJ','LZW','HDWSKMJW','GX','ESCAFY','S','VSAKQUZSAF','OGMDV','TW','OGJLZ','LZW','LJGMTDW','GX','YWLLAFY','MH','SFV','HAUCAFY','LZW','VSAKAWK','OZWF','KMVVWFDQ','S','OZALW','JSTTAL','OALZ','HAFC','WQWK','JSF','UDGKW','TQ','ZWJ','LZWJW','OSK','FGLZAFY','KG','NWJQ','JWESJCSTDW','AF','LZSL','FGJ','VAV','SDAUW','LZAFC','AL','KG','NWJQ','EMUZ','GML','GX','LZW','OSQ','LG','ZWSJ','LZW','JSTTAL','KSQ','LG','ALKWDX','GZ','VWSJ','GZ','VWSJ','A','KZSDD','TW','DSLW','OZWF','KZW','LZGMYZL','AL','GNWJ','SXLWJOSJVK','AL','GUUMJJWV','LG','ZWJ','LZSL','KZW','GMYZL','LG','ZSNW','OGFVWJWV','SL','LZAK','TML','SL','LZW','LAEW','AL','SDD','KWWEWV','IMALW','FSLMJSD','TML','OZWF','LZW','JSTTAL','SULMSDDQ','LGGC','S','OSLUZ','GML','GX','ALK','OSAKLUGSL','HGUCWL','SFV','DGGCWV','SL','AL','SFV','LZWF','ZMJJAWV','GF','SDAUW','KLSJLWV','LG','ZWJ','XWWL','XGJ','AL','XDSKZWV','SUJGKK','ZWJ','EAFV','LZSL','KZW','ZSV','FWNWJ','TWXGJW','KWWF','S','JSTTAL','OALZ','WALZWJ','S','OSAKLUGSLHGUCWL','GJ','S','OSLUZ','LG','LSCW','GML','GX','AL','SFV','TMJFAFY','OALZ','UMJAGKALQ','KZW','JSF','SUJGKK','LZW','XAWDV','SXLWJ','AL','SFV','XGJLMFSLWDQ','OSK','BMKL','AF','LAEW','LG','KWW','AL','HGH','VGOF','S','DSJYW','JSTTALZGDW','MFVWJ','LZW','ZWVYW','AF','SFGLZWJ','EGEWFL','VGOF','OWFL','SDAUW','SXLWJ','AL','FWNWJ','GFUW','UGFKAVWJAFY','ZGO','AF','LZW','OGJDV','KZW','OSK','LG','YWL','GML','SYSAF','LZW','JSTTALZGDW','OWFL','KLJSAYZL','GF','DACW','S','LMFFWD','XGJ','KGEW','OSQ','SFV','LZWF','VAHHWV','KMVVWFDQ','VGOF','KG','KMVVWFDQ','LZSL','SDAUW','ZSV','FGL','S','EGEWFL','LG','LZAFC','STGML','KLGHHAFY','ZWJKWDX','TWXGJW','KZW','XGMFV','ZWJKWDX','XSDDAFY','VGOF','S','NWJQ','VWWH','OWDD','WALZWJ','LZW','OWDD','OSK','NWJQ','VWWH','GJ','KZW','XWDD','NWJQ','KDGODQ','XGJ','KZW','ZSV','HDWFLQ','GX','LAEW','SK','KZW','OWFL','VGOF','LG','DGGC','STGML','ZWJ','SFV','LG','OGFVWJ','OZSL','OSK','YGAFY','LG','ZSHHWF','FWPL','XAJKL','KZW','LJAWV','LG','DGGC','VGOF','SFV','ESCW','GML','OZSL','KZW','OSK','UGEAFY','LG','TML','AL','OSK','LGG','VSJC','LG','KWW','SFQLZAFY','LZWF','KZW','DGGCWV','SL','LZW','KAVWK','GX','LZW','OWDD','SFV','FGLAUWV','LZSL','LZWQ','OWJW','XADDWV','OALZ','UMHTGSJVK','SFV','TGGCKZWDNWK','ZWJW','SFV','LZWJW','KZW','KSO','ESHK','SFV','HAULMJWK','ZMFY','MHGF','HWYK','KZW','LGGC','VGOF','S','BSJ','XJGE','GFW','GX','LZW','KZWDNWK','SK','KZW','HSKKWV','AL','OSK','DSTWDDWV','GJSFYW','ESJESDSVW','TML','LG','ZWJ','YJWSL','VAKSHHGAFLEWFL','AL','OSK','WEHLQ','KZW','VAV','FGL','DACW','LG','VJGH','LZW','BSJ','XGJ','XWSJ','GX','CADDAFY','KGEWTGVQ','KG','ESFSYWV','LG','HML','AL','AFLG','GFW','GX','LZW','UMHTGSJVK','SK','KZW','XWDD','HSKL','AL','OWDD','LZGMYZL','SDAUW','LG','ZWJKWDX','SXLWJ','KMUZ','S','XSDD','SK','LZAK','A','KZSDD','LZAFC','FGLZAFY','GX','LMETDAFY','VGOF','KLSAJK','ZGO','TJSNW','LZWQDD','SDD','LZAFC','EW','SL','ZGEW','OZQ','A','OGMDVFL','KSQ','SFQLZAFY','STGML','AL','WNWF','AX','A','XWDD','GXX','LZW','LGH','GX','LZW','ZGMKW','OZAUZ','OSK','NWJQ','DACWDQ','LJMW','VGOF','VGOF','VGOF','OGMDV','LZW','XSDD','FWNWJ','UGEW','LG','SF','WFV','A','OGFVWJ','ZGO','ESFQ','EADWK','ANW','XSDDWF','TQ','LZAK','LAEW','KZW','KSAV','SDGMV','A','EMKL','TW','YWLLAFY','KGEWOZWJW','FWSJ','LZW','UWFLJW','GX','LZW','WSJLZ','DWL','EW','KWW','LZSL','OGMDV','TW','XGMJ','LZGMKSFV','EADWK','VGOF','A','LZAFC','XGJ','QGM','KWW','SDAUW','ZSV','DWSJFL','KWNWJSD','LZAFYK','GX','LZAK','KGJL','AF','ZWJ','DWKKGFK','AF','LZW','KUZGGDJGGE','SFV','LZGMYZ','LZAK','OSK','FGL','S','NWJQ','YGGV','GHHGJLMFALQ','XGJ','KZGOAFY','GXX','ZWJ','CFGODWVYW','SK','LZWJW','OSK','FG','GFW','LG','DAKLWF','LG','ZWJ','KLADD','AL','OSK','YGGV','HJSULAUW','LG','KSQ','AL','GNWJ','QWK','LZSLK','STGML','LZW','JAYZL','VAKLSFUWTML','LZWF','A','OGFVWJ','OZSL','DSLALMVW','GJ','DGFYALMVW','ANW','YGL','LG','SDAUW','ZSV','FG','AVWS','OZSL','DSLALMVW','OSK','GJ','DGFYALMVW','WALZWJ','TML','LZGMYZL','LZWQ','OWJW','FAUW','YJSFV','OGJVK','LG','KSQ','HJWKWFLDQ','KZW','TWYSF','SYSAF','A','OGFVWJ','AX','A','KZSDD','XSDD','JAYZL','LZJGMYZ','LZW','WSJLZ','ZGO','XMFFQ','ALDD','KWWE','LG','UGEW','GML','SEGFY','LZW','HWGHDW','LZSL','OSDC','OALZ','LZWAJ','ZWSVK','VGOFOSJV','LZW','SFLAHSLZAWK','A','LZAFC','KZW','OSK','JSLZWJ','YDSV','LZWJW','OSK','FG','GFW','DAKLWFAFY','LZAK','LAEW','SK','AL','VAVFL','KGMFV','SL','SDD','LZW','JAYZL','OGJV','TML','A','KZSDD','ZSNW','LG','SKC','LZWE','OZSL','LZW','FSEW','GX','LZW','UGMFLJQ','AK','QGM','CFGO','HDWSKW','ESSE','AK','LZAK','FWO','RWSDSFV','GJ','SMKLJSDAS','SFV','KZW','LJAWV','LG','UMJLKWQ','SK','KZW','KHGCWXSFUQ','UMJLKWQAFY','SK','QGMJW','XSDDAFY','LZJGMYZ','LZW','SAJ','VG','QGM','LZAFC','QGM','UGMDV','ESFSYW','AL','SFV','OZSL','SF','AYFGJSFL','DALLDW','YAJD','KZWDD','LZAFC','EW','XGJ','SKCAFY','FG','ALDD','FWNWJ','VG','LG','SKC','HWJZSHK','A','KZSDD','KWW','AL','OJALLWF','MH','KGEWOZWJW','VGOF','VGOF','VGOF','LZWJW','OSK','FGLZAFY','WDKW','LG','VG','KG','SDAUW','KGGF','TWYSF','LSDCAFY','SYSAF','VAFSZDD','EAKK','EW','NWJQ','EMUZ','LGFAYZL','A','KZGMDV','LZAFC','VAFSZ','OSK','LZW','USL','A','ZGHW','LZWQDD','JWEWETWJ','ZWJ','KSMUWJ','GX','EADC','SL','LWSLAEW','VAFSZ','EQ','VWSJ','A','OAKZ','QGM','OWJW','VGOF','ZWJW','OALZ','EW','LZWJW','SJW','FG','EAUW','AF','LZW','SAJ','AE','SXJSAV','TML','QGM','EAYZL','USLUZ','S','TSL','SFV','LZSLK','NWJQ','DACW','S','EGMKW','QGM','CFGO','TML','VG','USLK','WSL','TSLK','A','OGFVWJ','SFV','ZWJW','SDAUW','TWYSF','LG','YWL','JSLZWJ','KDWWHQ','SFV','OWFL','GF','KSQAFY','LG','ZWJKWDX','AF','S','VJWSEQ','KGJL','GX','OSQ','VG','USLK','WSL','TSLK','VG','USLK','WSL','TSLK','SFV','KGEWLAEWK','VG','TSLK','WSL','USLK','XGJ','QGM','KWW','SK','KZW','UGMDVFL','SFKOWJ','WALZWJ','IMWKLAGF','AL','VAVFL','EMUZ','ESLLWJ','OZAUZ','OSQ','KZW','HML','AL','KZW','XWDL','LZSL','KZW','OSK','VGRAFY','GXX','SFV','ZSV','BMKL','TWYMF','LG','VJWSE','LZSL','KZW','OSK','OSDCAFY','ZSFV','AF','ZSFV','OALZ','VAFSZ','SFV','KSQAFY','LG','ZWJ','NWJQ','WSJFWKLDQ','FGO','VAFSZ','LWDD','EW','LZW','LJMLZ','VAV','QGM','WNWJ','WSL','S','TSL','OZWF','KMVVWFDQ','LZMEH','LZMEH','VGOF','KZW','USEW','MHGF','S','ZWSH','GX','KLAUCK','SFV','VJQ','DWSNWK','SFV','LZW','XSDD','OSK','GNWJ','SDAUW','OSK','FGL','S','TAL','ZMJL','SFV','KZW','BMEHWV','MH','GF','LG','ZWJ','XWWL','AF','S','EGEWFL','KZW','DGGCWV','MH','TML','AL','OSK','SDD','VSJC','GNWJZWSV','TWXGJW','ZWJ','OSK','SFGLZWJ','DGFY','HSKKSYW','SFV','LZW','OZALW','JSTTAL','OSK','KLADD','AF','KAYZL','ZMJJQAFY','VGOF','AL','LZWJW','OSK','FGL','S','EGEWFL','LG','TW','DGKL','SOSQ','OWFL','SDAUW','DACW','LZW','OAFV','SFV','OSK','BMKL','AF','LAEW','LG','ZWSJ','AL','KSQ','SK','AL','LMJFWV','S','UGJFWJ','GZ','EQ','WSJK','SFV','OZAKCWJK','ZGO','DSLW','ALK','YWLLAFY','KZW','OSK','UDGKW','TWZAFV','AL','OZWF','KZW','LMJFWV','LZW','UGJFWJ','TML','LZW','JSTTAL','OSK','FG','DGFYWJ','LG','TW','KWWF','KZW','XGMFV','ZWJKWDX','AF','S','DGFY','DGO','ZSDD','OZAUZ','OSK','DAL','MH','TQ','S','JGO','GX','DSEHK','ZSFYAFY','XJGE','LZW','JGGX','LZWJW','OWJW','VGGJK','SDD','JGMFV','LZW','ZSDD','TML','LZWQ','OWJW','SDD','DGUCWV','SFV','OZWF','SDAUW','ZSV','TWWF','SDD','LZW','OSQ','VGOF','GFW','KAVW','SFV','MH','LZW','GLZWJ','LJQAFY','WNWJQ','VGGJ','KZW','OSDCWV','KSVDQ','VGOF','LZW','EAVVDW','OGFVWJAFY','ZGO','KZW','OSK','WNWJ','LG','YWL','GML','SYSAF','KMVVWFDQ','KZW','USEW','MHGF','S','DALLDW','LZJWWDWYYWV','LSTDW','SDD','ESVW','GX','KGDAV','YDSKK','LZWJW','OSK','FGLZAFY','GF','AL','WPUWHL','S','LAFQ','YGDVWF','CWQ','SFV','SDAUWK','XAJKL','LZGMYZL','OSK','LZSL','AL','EAYZL','TWDGFY','LG','GFW','GX','LZW','VGGJK','GX','LZW','ZSDD','TML','SDSK','WALZWJ','LZW','DGUCK','OWJW','LGG','DSJYW','GJ','LZW','CWQ','OSK','LGG','KESDD','TML','SL','SFQ','JSLW','AL','OGMDV','FGL','GHWF','SFQ','GX','LZWE','ZGOWNWJ','GF','LZW','KWUGFV','LAEW','JGMFV','KZW','USEW','MHGF','S','DGO','UMJLSAF','KZW','ZSV','FGL','FGLAUWV','TWXGJW','SFV','TWZAFV','AL','OSK','S','DALLDW','VGGJ','STGML','XAXLWWF','AFUZWK','ZAYZ','KZW','LJAWV','LZW','DALLDW','YGDVWF','CWQ','AF','LZW','DGUC','SFV','LG','ZWJ','YJWSL','VWDAYZL','AL','XALLWV','SDAUW','GHWFWV','LZW','VGGJ','SFV','XGMFV','LZSL','AL','DWV','AFLG','S','KESDD','HSKKSYW','FGL','EMUZ','DSJYWJ','LZSF','S','JSLZGDW','KZW','CFWDL','VGOF','SFV','DGGCWV','SDGFY','LZW','HSKKSYW','AFLG','LZW','DGNWDAWKL','YSJVWF','QGM','WNWJ','KSO','ZGO','KZW','DGFYWV','LG','YWL','GML','GX','LZSL','VSJC','ZSDD','SFV','OSFVWJ','STGML','SEGFY','LZGKW','TWVK','GX','TJAYZL','XDGOWJK','SFV','LZGKW','UGGD','XGMFLSAFK','TML','KZW','UGMDV','FGL','WNWF','YWL','ZWJ','ZWSV','LZJGMYZ','LZW','VGGJOSQ','SFV','WNWF','AX','EQ','ZWSV','OGMDV','YG','LZJGMYZ','LZGMYZL','HGGJ','SDAUW','AL','OGMDV','TW','GX','NWJQ','DALLDW','MKW','OALZGML','EQ','KZGMDVWJK','GZ','ZGO','A','OAKZ','A','UGMDV','KZML','MH','DACW','S','LWDWKUGHW','A','LZAFC','A','UGMDV','AX','A','GFDQ','CFGO','ZGO','LG','TWYAF','XGJ','QGM','KWW','KG','ESFQ','GMLGXLZWOSQ','LZAFYK','ZSV','ZSHHWFWV','DSLWDQ','LZSL','SDAUW','ZSV','TWYMF','LG','LZAFC','LZSL','NWJQ','XWO','LZAFYK','AFVWWV','OWJW','JWSDDQ','AEHGKKATDW','LZWJW','KWWEWV','LG','TW','FG','MKW','AF','OSALAFY','TQ','LZW','DALLDW','VGGJ','KG','KZW','OWFL','TSUC','LG','LZW','LSTDW','ZSDX','ZGHAFY','KZW','EAYZL','XAFV','SFGLZWJ','CWQ','GF','AL','GJ','SL','SFQ','JSLW','S','TGGC','GX','JMDWK','XGJ','KZMLLAFY','HWGHDW','MH','DACW','LWDWKUGHWK','LZAK','LAEW','KZW','XGMFV','S','DALLDW','TGLLDW','GF','AL','OZAUZ','UWJLSAFDQ','OSK','FGL','ZWJW','TWXGJW','KSAV','SDAUW','SFV','JGMFV','LZW','FWUC','GX','LZW','TGLLDW','OSK','S','HSHWJ','DSTWD','OALZ','LZW','OGJVK','VJAFC','EW','TWSMLAXMDDQ','HJAFLWV','GF','AL','AF','DSJYW','DWLLWJK','AL','OSK','SDD','NWJQ','OWDD','LG','KSQ','VJAFC','EW','TML','LZW','OAKW','DALLDW','SDAUW','OSK','FGL','YGAFY','LG','VG','LZSL','AF','S','ZMJJQ','FG','ADD','DGGC','XAJKL','KZW','KSAV','SFV','KWW','OZWLZWJ','ALK','ESJCWV','HGAKGF','GJ','FGL','XGJ','KZW','ZSV','JWSV','KWNWJSD','FAUW','DALLDW','ZAKLGJAWK','STGML','UZADVJWF','OZG','ZSV','YGL','TMJFL','SFV','WSLWF','MH','TQ','OADV','TWSKLK','SFV','GLZWJ','MFHDWSKSFL','LZAFYK','SDD','TWUSMKW','LZWQ','OGMDV','FGL','JWEWETWJ','LZW','KAEHDW','JMDWK','LZWAJ','XJAWFVK','ZSV','LSMYZL','LZWE','KMUZ','SK','LZSL','S','JWVZGL','HGCWJ','OADD','TMJF','QGM','AX','QGM','ZGDV','AL','LGG','DGFY','SFV','LZSL','AX','QGM','UML','QGMJ','XAFYWJ','NWJQ','VWWHDQ','OALZ','S','CFAXW','AL','MKMSDDQ','TDWWVK','SFV','KZW','ZSV','FWNWJ','XGJYGLLWF','LZSL','AX','QGM','VJAFC','EMUZ','XJGE','S','TGLLDW','ESJCWV','HGAKGF','AL','AK','SDEGKL','UWJLSAF','LG','VAKSYJWW','OALZ','QGM','KGGFWJ','GJ','DSLWJ','ZGOWNWJ','LZAK','TGLLDW','OSK','FGL','ESJCWV','HGAKGF','KG','SDAUW','NWFLMJWV','LG','LSKLW','AL','SFV','XAFVAFY','AL','NWJQ','FAUW','AL','ZSV','AF','XSUL','S','KGJL','GX','EAPWV','XDSNGMJ','GX','UZWJJQLSJL','UMKLSJV','HAFWSHHDW','JGSKL','LMJCWQ','LGXXWW','SFV','ZGL','TMLLWJWV','LGSKL','KZW','NWJQ','KGGF','XAFAKZWV','AL','GXX','OZSL','S','UMJAGMK','XWWDAFY','KSAV','SDAUW','A','EMKL','TW','KZMLLAFY','MH','DACW','S','LWDWKUGHW','SFV','KG','AL','OSK','AFVWWV','KZW','OSK','FGO','GFDQ','LWF','AFUZWK','ZAYZ','SFV','ZWJ','XSUW','TJAYZLWFWV','MH','SL','LZW','LZGMYZL','LZSL','KZW','OSK','FGO','LZW','JAYZL','KARW','XGJ','YGAFY','LZJGMYZ','LZW','DALLDW','VGGJ','AFLG','LZSL','DGNWDQ','YSJVWF','XAJKL','ZGOWNWJ','KZW','OSALWV','XGJ','S','XWO','EAFMLWK','LG','KWW','AX','KZW','OSK','YGAFY','LG','KZJAFC','SFQ','XMJLZWJ','KZW','XWDL','S','DALLDW','FWJNGMK','STGML','LZAK','XGJ','AL','EAYZL','WFV','QGM','CFGO','KSAV','SDAUW','LG','ZWJKWDX','AF','EQ','YGAFY','GML','SDLGYWLZWJ','DACW','S','USFVDW','A','OGFVWJ','OZSL','A','KZGMDV','TW','DACW','LZWF','SFV','KZW','LJAWV','LG','XSFUQ','OZSL','LZW','XDSEW','GX','S','USFVDW','AK','DACW','SXLWJ','LZW','USFVDW','AK','TDGOF','GML','XGJ','KZW','UGMDV','FGL','JWEWETWJ','WNWJ','ZSNAFY','KWWF','KMUZ','S','LZAFY','SXLWJ','S','OZADW','XAFVAFY','LZSL','FGLZAFY','EGJW','ZSHHWFWV','KZW','VWUAVWV','GF','YGAFY','AFLG','LZW','YSJVWF','SL','GFUW','TML','SDSK','XGJ','HGGJ','SDAUW','OZWF','KZW','YGL','LG','LZW','VGGJ','KZW','XGMFV','KZW','ZSV','XGJYGLLWF','LZW','DALLDW','YGDVWF','CWQ','SFV','OZWF','KZW','OWFL','TSUC','LG','LZW','LSTDW','XGJ','AL','KZW','XGMFV','KZW','UGMDV','FGL','HGKKATDQ','JWSUZ','AL','KZW','UGMDV','KWW','AL','IMALW','HDSAFDQ','LZJGMYZ','LZW','YDSKK','SFV','KZW','LJAWV','ZWJ','TWKL','LG','UDAET','MH','GFW','GX','LZW','DWYK','GX','LZW','LSTDW','TML','AL','OSK','LGG','KDAHHWJQ','SFV','OZWF','KZW','ZSV','LAJWV','ZWJKWDX','GML','OALZ','LJQAFY','LZW','HGGJ','DALLDW','LZAFY','KSL','VGOF','SFV','UJAWV','UGEW','LZWJWK','FG','MKW','AF','UJQAFY','DACW','LZSL','KSAV','SDAUW','LG','ZWJKWDX','JSLZWJ','KZSJHDQ','A','SVNAKW','QGM','LG','DWSNW','GXX','LZAK','EAFMLW','KZW','YWFWJSDDQ','YSNW','ZWJKWDX','NWJQ','YGGV','SVNAUW','LZGMYZ','KZW','NWJQ','KWDVGE','XGDDGOWV','AL','SFV','KGEWLAEWK','KZW','KUGDVWV','ZWJKWDX','KG','KWNWJWDQ','SK','LG','TJAFY','LWSJK','AFLG','ZWJ','WQWK','SFV','GFUW','KZW','JWEWETWJWV','LJQAFY','LG','TGP','ZWJ','GOF','WSJK','XGJ','ZSNAFY','UZWSLWV','ZWJKWDX','AF','S','YSEW','GX','UJGIMWL','KZW','OSK','HDSQAFY','SYSAFKL','ZWJKWDX','XGJ','LZAK','UMJAGMK','UZADV','OSK','NWJQ','XGFV','GX','HJWLWFVAFY','LG','TW','LOG','HWGHDW','TML','ALK','FG','MKW','FGO','LZGMYZL','HGGJ','SDAUW','LG','HJWLWFV','LG','TW','LOG','HWGHDW','OZQ','LZWJWK','ZSJVDQ','WFGMYZ','GX','EW','DWXL','LG','ESCW','GFW','JWKHWULSTDW','HWJKGF','KGGF','ZWJ','WQW','XWDD','GF','S','DALLDW','YDSKK','TGP','LZSL','OSK','DQAFY','MFVWJ','LZW','LSTDW','KZW','GHWFWV','AL','SFV','XGMFV','AF','AL','S','NWJQ','KESDD','USCW','GF','OZAUZ','LZW','OGJVK','WSL','EW','OWJW','TWSMLAXMDDQ','ESJCWV','AF','UMJJSFLK','OWDD','ADD','WSL','AL','KSAV','SDAUW','SFV','AX','AL','ESCWK','EW','YJGO','DSJYWJ','A','USF','JWSUZ','LZW','CWQ','SFV','AX','AL','ESCWK','EW','YJGO','KESDDWJ','A','USF','UJWWH','MFVWJ','LZW','VGGJ','KG','WALZWJ','OSQ','ADD','YWL','AFLG','LZW','YSJVWF','SFV','A','VGFL','USJW','OZAUZ','ZSHHWFK','KZW','SLW','S','DALLDW','TAL','SFV','KSAV','SFPAGMKDQ','LG','ZWJKWDX','OZAUZ','OSQ','OZAUZ','OSQ','ZGDVAFY','ZWJ','ZSFV','GF','LZW','LGH','GX','ZWJ','ZWSV','LG','XWWD','OZAUZ','OSQ','AL','OSK','YJGOAFY','SFV','KZW','OSK','IMALW','KMJHJAKWV','LG','XAFV','LZSL','KZW','JWESAFWV','LZW','KSEW','KARW','LG','TW','KMJW','LZAK','YWFWJSDDQ','ZSHHWFK','OZWF','GFW','WSLK','USCW','TML','SDAUW','ZSV','YGL','KG','EMUZ','AFLG','LZW','OSQ','GX','WPHWULAFY','FGLZAFY','TML','GMLGXLZWOSQ','LZAFYK','LG','ZSHHWF','LZSL','AL','KWWEWV','IMALW','VMDD','SFV','KLMHAV','XGJ','DAXW','LG','YG','GF','AF','LZW','UGEEGF','OSQ','KG','KZW','KWL','LG','OGJC','SFV','NWJQ','KGGF','XAFAKZWV','GXX','LZW','USCW']

df = createVis(newMessage, shiftValue)
assert( df[ df['letter'] == 'T' ]['count'].max() == 867 ), "Are you sure you found the correct solution?"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()
print("Here's the story (from Alice in Wonderland):")
s = decodeGivenShift(newMessage, shiftValue)
print(' '.join(s))
```

## Submit Your Work!

You're almost done -- congratulations!  

You need to do a few more things:

1. Save your work.  To do this, create a **notebook checkpoint** by using the menu within the notebook to go **File -> Save and Checkpoint**


2. After you have saved and checkpointed, exit this notebook by going to **File -> Close and Halt**


3. Once you have closed your notebook select **Quit** on the **Jupyter main page** to stop the Jupyter server


4. Return to the Data Science Discovery page on how to use git to turn this notebook into the course!