# lab_pandas

Hi Data Scientists,

Welcome to the second lab!  In this lab, you will explore **data types** in Python and learn the first tool of Data Science -- `pandas`.  As you go through the lab:

- Make sure to run every cell with Python code and observe the output
- Complete all puzzles and submit this lab before Tuesday evening at 11:59pm

### Labs are Collabrative

In your lab section (right now!), find a person or two around you and form a group.  Work on this lab together, making sure all of you understand each component of the lab.

### Part 0: Your Group!

Edit the next Python cell to add information about who you're working with in your lab section:

```python
# First, meet your CAs and TA if you haven't already!
# ...first name is enough, we'll know who they are! :)
ta_name = ""
ca1_name = ""
ca2_name = ""


# Also, make sure to meet your team for this lab!  Find out their name, what major they're in,
# and learn something new about htem that you never knew more!
partner1_name = ""
partner1_netid = ""
partner1_major = ""

partner2_name = ""
partner2_netid = ""
partner2_major = ""

partner3_name = ""
partner3_netid = ""
partner3_major = ""
```

### Getting Help

Remember, there are a lot of ways to get help if you find yourself stuck:

1. In lab section, your TA, CAs, and your peers are here for you!

2. On the course Piazza page: https://piazza.com/illinois/spring2020/stat107

3. Office Hours:

  - Open office hours **every Monday, Tuesday, Wednesday, and Friday** from 4:00pm - 6:00pm in 23 Illini Hall

  - Prof. Wade + Prof. Karle are available **every Wednesday** from 8:30am - 10:00am in 2215 Siebel Center

## Part 1: Data Types

In Python, data must be stored as a **number** (without quotes), as a **string** (text with quotes), or as some **other type**.  Every piece of data in Python will have three **features**:

1. A **name** *(this is what we refer to it by)*
2. A **value** *(this is the value is stores)*
3. A **data type** *(this is the format that the data is stored)*

### Example: Data Types

Run the following line that creates the name `seconds_in_a_minute`:

```python
seconds_in_a_minute = 60 * 60
seconds_in_a_minute
```

The features of `seconds_in_a_minute` are the following:

- The **name** is `seconds_in_a_minute`
- The **value** of `seconds_in_a_minute` is `3600`
- The **data type** of `seconds_in_a_minute` is a **number**

### Puzzle 1.1: Your Name
#### Puzzle 1.1a: Creating Your Name

Modify the following cell of Python to replace `...` with your name:

- Remember, because your name is a **string**, you will need to use quotes (`'`...`'`)

<!-- grade_id: cell-522cdea9b9b77bb1, points: 0, grade: false, solution: true -->
```python
my_name = ...
my_name

```

<!-- grade_id: cell-db03062e5907fe0f, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1.1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(type(my_name) == type('string')), 'It looks like your name is not a string.  Make sure to assign your name.'
```

#### Puzzle 1.1b: Features of `my_name`

Modify the following cell of to replace ... with the correct features about `my_name`.  (This is a text cell, so we won't see the results until we grade it later.)

- The **name** is ...
- The **value** of `my_name` is ...
- The **data type** of `my_name` is ...

### Puzzle 1.2: The Current Year

#### Puzzle 1.2a: Creating `current_year`

Modify the following cell of Python to replace `...` with the current year (as a **number**):

<!-- grade_id: cell-d0842fa179dbcea7, points: 0, grade: false, solution: true -->
```python
current_year = ...
current_year

```

<!-- grade_id: cell-a80420d1ef2b0104, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1.2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(type(current_year) == type(2020)), 'It looks like your name is not a number.  Make sure to assign the year as a number.'
```

#### Puzzle 1.2b: Features of `current_year`

Modify the following cell of to replace ... with the correct features about `current_year`:

- The **name** is ...
- The **value** of `current_year` is ...
- The **data type** of `current_year` is ...

### Review: Operations on Numbers

Last week, you used the **add** (`+`) operation to add two **numbers** together.  For example:

```python
eleven = 4 + 7
eleven
```

You also can use **subtraction** (`-`), **multiplication** (`*`), **division** (`/`), and you can **raise a number to a power** (`**`):

```python
fiftyfive = 100 - 45
print(fiftyfive)

sixteen = 8 * 2
print(sixteen)

two = 6 / 3
print(two)

eight = 2 ** 3
print(eight)
```

A few things you may have noticed:
- We used `print()` again this week, allowing us to output values that are not the last line of code! 
- The result of `two` is `2.0`, a non-whole number!  There are two different types of numbers: **integers** (whole numbers) and **floating point numbers** (numbers with a decimal).  We will explore these in the next part of this lab!

### Concatenating Strings

In Python, you join two **strings** together with the same `+` operation you used with **numbers**.  For example, we can join `'Hello'` and `'World'` together:

```python
print('Hello' + 'World')
```

...that's a little awkward.  Let's add a space together to join three strings together:

```python
print('Hello' + ' ' + 'World')
```

When we join strings together, we do **NOT** say that we are **adding** strings together.  We **add** numbers together.  Instead, this operation of joining two or more strings is called **concatenation** *(concatenation literally means "the action of linking things together in a series")*.

### Puzzle 1.3: 

Instead of `'Hello World'`, modify the value in `hello` to say hello **to you** by using `my_name` you created earlier in this lab:

<!-- grade_id: cell-47e9c78f29292a6a, points: 0, grade: false, solution: true -->
```python
# Modify this Python to say 'Hello ___', where ___ is your name.
# - Use `my_name` you created earlier.

hello = 'Hello' + ' ' + 'World'
print(hello)

```

<!-- grade_id: cell-13f3d582c7941841, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1.3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(hello == 'Hello ' + my_name), "Your string `hello` is not saying hi to yourself. :("
```

### Puzzle 1.4:

Let's wish the current year (2020) hello as well!  Modify the following code to print hello **to the current year** using `current_year` you created earlier in this lab:

- If you do this corretly, **you will see an error** -- an error is the expected result!!

```python
# Modify this Python to say "Hello 2020", where 2020 is the current year.
# - Use `current_year` you created earlier.
print('Hello' + ' ' + 'World')
```

If you used `current_year`, Python will be confused and give you the following error:

![TypeError: can only concatenate str to str](img/error-cannot-concat.png)

The text of the error tells us something about Python:

- Python can only concatenate a **string** to another **string**
- Therefore, Python can **NOT** concatenate a **string** to a **number**

In the next section, we will learn how to convert between data types!


## Part 2: Functions: Programs that work for us!

In programming, we do not want to do everything ourselves.  Instead, **functions** provide us **functionality** that does work for us!  We already have learned one **function**:

- `print( value )`, prints out `value` to the screen

All **functions** always are made up of three parts:

1. Every function always has a **name** (ex: `print`)
2. Every function's name is always following with **paired parentheses** (ex: `print()`)
3. Functions may take **values between the paired parentheses** (ex: `print('Hello')` or `print(43)`).  We call the values that a function has between it's paired parentheses the **function parameters**.

Here's a visual breakdown of this same information:

![Visual print function](img/print-function.png)

The values that a function can receive vary by function and must be known.  `print` is very special in that it can take in almost any type of data.

### The `str` Function

The second function we will use is the `str` function.  This function has a different name than `print` and will do something different.  However, we use `str` is the same way.

The `str` function will take in a value as it's parameter and **return** a string version of that value.  For example, the **number** `42` can be represented as the string `"42"`.  The **return value** of the function is what is left after the function is run.

A few examples to run, notice the output of each of them:

```python
str(42)
```

```python
str('Hello')
```

```python
nine = str(4 + 5)
nine
```

### Puzzle 2.1:

Now we can wish the current year (2020) hello by using the `str` function!

Modify the following code to say hello **to the current year** using `current_year` you created back in Part 1 of this lab:

```python
# Modify this Python to say "Hello 2020", where 2020 is the current year.
# - Use `current_year` you created earlier.
print('Hello' + ' ' + 'World')
```

### The `float` and `int` functions

The next two functions do the reverse of `str`.  Instead of taking a value and returning a string, `float` and `int` are functions that take a value and return a number.  Both `float` and `int` are similar, but have an important difference:

- `int` (short for **integer**) returns whole numbers only (ex: 1, 2, 3, 4, etc)
- `float` (short for **floating point**) returns decimal numbers (ex: 5.2, 3.141592654, 4.25, 11.0, etc)

If you use `int` on a non-whole number, `int` will return only the whole number component of the number.  For example, the whole number part of **3.9** is **3** (note that it does **NOT** do any rounding).  A few examples:

```python
print( int(4.9) )
print( int(-8.7) )
print( float("4.8") )
print( int("9") )
```

If Python is unable to convert the value to a number, your code will have an error:

```python
float('Hello')
```

### Puzzle 2.2:

In the next cell, several names are defined with various values.  Run this cell:

```python
a = "17"
b = "44.2"
c = "19"
```

**Without editing the previous cell**, compute the value of $a + b + c$ in the next cell:

<!-- grade_id: cell-53d3714573ab7c3e, points: 0, grade: false, solution: true -->
```python
# Compute the value of a + b + c
total = ...
total

```

Note: Your result should **NOT** be `1744.219`.  Make sure your answer makes sense before running the test cases.

<!-- grade_id: cell-e99fd2370918408c, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2.2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(a == "17"), "You must not modify the value of `a`."
assert(type(a) == type("17")), "You must not modify the value of `a`."

assert(b == "44.2"), "You must not modify the value of `b`."
assert(type(b) == type("44.2")), "You must not modify the value of `b`."

assert(c == "19"), "You must not modify the value of `c`."
assert(type(c) == type("19")), "You must not modify the value of `c`."

assert(total == 80.2), "The value in `total` is not the sum of `a`, `b`, and `c`."
```

## Part 3: Importing Datasets

Data Science is focused on using tools to access, analyze, and find meaningful results in data.  Now that we know some basic Python programming, we are ready to do begin to use our first dataset!

### CSV File Format

One widely used file format to store datasets is as **CSV** or **Comma-Separated Values**.  The CSV file maps extremely well between a spreadsheet program (ex: Microsoft Excel, Google Sheets) and programming languages.  Specifically:

- Every row on a spreadsheet is a line in a CSV file
- Every column on a spreadsheet is separated by a comma in a CSV file

### Download the Illinois Course Catalog as a CSV

Let's grab the first dataset we will use!  For this dataset, we have already collected and put the data into a Google Sheet for you.  You will need to download the file as a CSV and place it in the same directory as this Python notebook.

1. Open the Google Sheet for the "Illinois Catalog Dataset": https://docs.google.com/spreadsheets/d/1ZnpZpipyui_zIIBblmalRYiyEG1glUKinqJtQxVL04o/edit#gid=0

2. Choose **File** -> **Download As** -> **Comma-separated Values (.csv, current sheet)**

3. Find the downloaded file on your computer and move it into the same directory as this notebook.
  * The file is likely in your Downloads as `Illinois Course Catalog - Sheet1.csv`
  * You notebook is likely on your Desktop under `stat107 -> [netid] -> lab_pandas`
  * You need to ensure the downloaded file is in the `lab_pandas` folder before running the next cell.

### Using Pandas to Read CSV Files

In Data Science DISCOVERY, we will use the `pandas` **library** for working with data.  Libraries are collections of functions and other objects that extend what Python can do by itself.  The `pandas` library is the most widely used library for data science.

To use the `pandas` library, we must **import** it into your notebook.  In future labs, you will find this will be one of the first lines we will run in every notebook:

```python
import pandas as pd
```

To use a function inside of the pandas library, the function name will always start with `pandas.` to denote that Python can find it within the `pandas` library.  The first function we will use from the `pandas` library is `read_csv`.
- Because `read_csv` is in the pandas library, the full function name is `pandas.read_csv()`.

### Puzzle 3.1:

The `pd.read_csv` function's first parameter is the name of the file that it is reading.  Using this, we know the format of the function follows the expected format.  Set the name `df` to be the result of the `pd.read_csv` function:

<!-- grade_id: cell-50e94d6dd49949e2, points: 0, grade: false, solution: true -->
```python
df = ...
df

```

<!-- grade_id: cell-8816926a2c95fc6f, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3.1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert( len(df) == 8589 ), "It looks like you did not import the correct dataset?"
```

### The DataFrame Type

In Python so far, we have had **string** and **number** data types.  With the successful `pd.read_csv` function, we now have a third type: a **DataFrame**.  We can always ask Python to tell us the type with the `type` function:

```python
type(df)
```

The `pandas.core.frame.DataFrame` (which we will just call `DataFrame`) type has hundreds of useful functions that can be called on it:

- We can choose data that matches a specific requirement,
- We can group our data by a certain column,
- ...and hundreds of other things...

The key idea here is that **`df` is the name of the DataFrame** that contains the data in the CSV file you downloaded.

### Accessing a Column

The dataset you downloaded has six columns: `Year`, `Term`, `YearTerm`, `Subject`, `Number`, and `Title`.  We can access a column by **indexing into the DataFrame**.  Whenever we **index into** something, we will use matchingsquare brackets: `[` and `]`.  For example, to **index into the column `Subject`**:

```python
df["Subject"]
```

Notice that we now only have the **Subject** data in our output, no longer the whole DataFrame.  We can do this with any of the six columns using the same format:

![image.png](img/DataFrame-index-into-column.png)

### Puzzle 3.2:

List the **name of every course** (ex: 'Data Science Discovery') at Illinois:

- You may need to refer back to the data sheet (or the DataFrame output above) to get the correct column name.

```python
...
```

### Finding a Subset of the Data

Let's find a subset of the data.  To do this, we will **index into the DataFrame** with a **conditional statement**.  A **conditional** is anything that Python can check if it's **True/Yes or False/No**.  For example:

- Is this course's subject STAT?  *...YES, that's True*
- Is this course offered in 2012?  *...NO, that's False*

#### Writing an Equality Conditional in Python

The most basic conditional is a test for equality -- if something is **exactly equal to** something else.  To test for equality, Python uses `==`.

- Remember, we already used `=` to set a name equal to something -- so Python needs to use something different.

In general, we will test if **each row in a column contains some value**.  For example, we may want to test if the subject for a given course is exactly equal to `'STAT'`.  To do this, the conditional will be:

![image.png](img/DataFrame-conditional.png)


### Puzzle 3.3: A standalone conditional

Write the conditional to check if each row in our DataFrame's **Subject** is `'STAT'`.

- You will expect to see a whole bunch of `False`.


```python
...
```

The expected result will list a large number or row numbers (0, 1, ..., 8741) that tells us if each row is STAT.  This is almost useful, but we want to **index the whole conditional into the DataFrame** to get **ONLY the rows where the conditional is True**.

We will index into the DataFrame with the full condiitional:

![image.png](img/DataFarme-row-conditional.png)

### Puzzle 3.4: Selecting a Subset of Rows

In the next cell, select only the rows that have a **Subject** of `'STAT'`:

```python
...
```

## Part 4: Exploring the DataSet

You now know the basics of `pandas`.  Let's explore the dataset!


### Puzzle 4.1: Selecting your Major's Rows

Select only the rows that have a **Subject** of your current/intended major.  Assign this to the value `mydf`:

```python
mydf = ...
mydf
```

```python
# There are no automatic test cases here -- you need to make sure your output is correct!
# Check over it to ensure it is what you expect to see!
```

### Puzzle 4.2: Finding your Upper-Level Courses

An "upper-level" course at Illinois is any and all courses with a number of `300` or greater.  Your major likely has many upper=level courses for you to take.

Using the same idea of a conditional, you will replace the logic of `==` (*which means "exactly equal to"*) with `>=` (*which means greater than or equal to*).  Write a conditional, using `mydf`, to find **all of the upper-level courses in your major** and store it in `mydf_upper_level`.

```python
mydf_upper_level = ...
mydf_upper_level
```

```python
# There are no automatic test cases here -- you need to make sure your output is correct!
# Check over it to ensure it is what you expect to see!
```

## Submit Your Work!

You're almost done -- congratulations!

You need to do two more things:

1. Save your work.  To do this, create a **notebook checkpoint** by using the menu within the notebook to go **File -> Save and Checkpoint**

2. Choose `File` and then `Close and Halt` from this notebook.

3. Choose `Quit` on the main notebook webpage.

4. Return to your command line and follow the directions on the Data Science Discovery page on how to use git to turn this notebook into the course!