# Central Limit Theorem (CLT) !

Welcome! In this lab, we will be exploring the Central Limit Theorem (CLT)! Although this concept was first developed by Abraham de Moivre in 1733, it wasn’t formally named until 1930, when noted Hungarian mathematician George Polya officially dubbed it the Central Limit Theorem.

## Review: What is the Central Limit Theorem?
Basically, the Central Limit Theorem says:  

If we repeatedly take independent random samples of size n from any population, then when n is large, the distribution of the sample means will approach a normal distribution.

Wow! This works for any population, even non-normal populations.  The CLT is a very useful theorem that is the foundation for many data science procedures. We'll see a lot of these procedures in the upcoming weeks in Data Science DISCOVERY!

Today, we will be working with a "song_data" dataset that has around 19,000 songs which have multiple attributes. This lab will help demonstrate how the Central Limit Theorem works! Below are the 15 attributes of the "song" dataset. 

The Attributes:

- `song_name`: contains the name of the song
- `song_popularity`: how popular a song is
- `song_duration_ms`: length of the song
- `acousticness`: the higher the value, the more acoustic the song is
- `danceability`: the higher the value, the easier it is to dance to this song
- `energy`: the energy of a song, the higher the value, the more energetic
- `instrumentalness`: the higher the value, the less vocals there are
- `key`: key the track is in
- `liveness`: the higher the value, the more of an audience presence there is
- `loudness`: the higher the value, the louder the song (in dB)
- `audiomode`: represents the modality of the track: major is represented by 1 and minor is represented by 0
- `speechiness`: the higher the value, the more presence of spoken words
- `tempo`: the overall estimated tempo of the song
- `time_signature`: an estimated overall time signature of a track
- `audio_valence`: the higher the value, the more positive it is

# 0. Import the Dataset

In the same folder as this notebook find the dataset `song_data.csv` that is provided. You know what to do to import the libraries and read the data! :)

<!-- grade_id: cell-615783f5a93483b3, points: 0, grade: false, solution: true -->
```python
df = ...


```

# 1. Puzzle 1 

### 1.1 Find Descriptive Statistics
Let's do some basic EDA to start. Find the mean, median, standard deviation, first quartile, and third quartile for each of the attributes. Quantiles are the same thing as quartiles that we talked about in class. *You will need some of these values later so make sure to store them as variables!*

<!-- grade_id: cell-1a9ba098db01bfaa, points: 0, grade: false, solution: true -->
```python
df_mean = ...
df_median = ...
df_sd = ...
df_quartile1st = ...
df_quartile3rd = ...



```

<!-- grade_id: cell-496495a67929f448, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1.1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(round(df_mean[0],0) == 53), "The value is not correct!"
assert(round(df_median[0],0) == 56), "The value is not correct!"
assert(round(df_sd[0],0) == 22), "The value is not correct!"
assert(round(df_quantile1st[0],0) == 40), "The value is not correct!"
assert(round(df_quantile3rd[0],0) == 69), "The value is not correct!"
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

### 1.2 Histograms for the Dataset

Plot the histograms for each column and find the two columns whose histograms are closest to the normal distribution. Calcuate the mean and standard deviation for these two columns. 

- If needed: In `df.hist`, use `figsize=(50, 30)` to make your figure bigger.


<!-- grade_id: cell-20bcf1a74c11c1ec, points: 0, grade: false, solution: true -->
```python
var1_mean = ...
var1_sd = ...
var2_mean = ...
var2_sd = ...



```

<!-- grade_id: cell-1abef90366d2a757, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1.2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(round(var1_mean,1) == 0.6), "The value is not correct!"
assert(round(var1_sd,1) == 0.2), "The value is not correct!"
assert(round(var2_mean,1) == 53), "The value is not correct!"
assert(round(var2_sd,1) == 21.9), "The value is not correct!"
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

# 2. Puzzle 2

### 2.1 Simulation #1
Now, we will see the beauty of the Central Limit Theorem.  

Define a funciton called `simulate_average` that will simulate playing `n` random songs and return the average value of these `n` songs for each column. For example, the column may be `'danceability'` and your function should return the average `'danceability'` for the `n` songs. (Hint: Your function should take 2 arguments)

<!-- grade_id: cell-6c5408ef21a469e8, points: 0, grade: false, solution: true -->
```python
### simulate_average ###



```

### 2.1.1 Approximating The Normal Distribution

Plot the histogram for the column `loudness`. Then using the `simulate_average` function we defined before, use the column `loudness` to run the simulation **100** times with `n = 1000`.  Store the results in a dataframe and plot the histogram to see the results.  Repeat this with the variable `danceability`.  Finally, repeat the process **10000** times for both variables and see the results.

<!-- grade_id: cell-de2def1a23d85b1a, points: 0, grade: false, solution: true -->
```python
loudness_hist = ...
df_danceability_100 = ...
df_loudness_100 = ...
df_danceability_10000 = ...
df_loudness_10000 = ...


```

<!-- grade_id: cell-fbb14166a355972b, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert( (0.62 <= df_danceability_100.mean()).bool() & (df_danceability_100.mean() <= 0.68).bool() ), "The simulation is not correct!"
assert( (-7.6 <= df_loudness_100.mean()).bool() & (df_loudness_100.mean() <= -7.3).bool() ), "The simulation is not correct!"
assert( (0.62 <= df_danceability_10000.mean()).bool() & (df_danceability_10000.mean() <= 0.68).bool() ), "The simulation is not correct!"
assert( (-7.6 <= df_loudness_10000.mean()).bool() & (df_loudness_10000.mean() <= -7.3).bool() ), "The simulation is not correct!"
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

### 2.1 Simulation #2
In simulation #1, we saw the CLT in action using the song dataset. In this simulation, we are going to see the CLT again, but for rolling dice. Create **10** simulations of **50** die rolls.  For each simulation, find the average of the dice rolls and store the value. Then create a histogram of your results. Repeat, but using **100** simulations. Finally, repeat again with **1000** simulations.

<!-- grade_id: cell-7c3bd1e8efe3344a, points: 0, grade: false, solution: true -->
```python
df_dice_10 = ...
df_dice_100 = ...
df_dice_1000 = ...



```

<!-- grade_id: cell-f850852667e6dce3, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert( (3.3 <= df_dice_1000.mean()).bool() & (df_dice_1000.mean() <= 3.6).bool() ), "The simulation is not correct!"
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

## Submit Your Work!

You're almost done -- congratulations!

You need to do two more things:

1. Save your work.  To do this, create a **notebook checkpoint** by using the menu within the notebook to go **File -> Save and Checkpoint**

2. Choose `File` and then `Close and Halt` from this notebook.

3. Choose `Quit` on the main notebook webpage.

4. Return to your command line and follow the directions on the Data Science Discovery page on how to use git to turn this notebook into the course!

```python

```