# Your Grades and GPA at Illinois!

In this lab, you will explore the GPA dataset and find out about the easiest, hardest, and most popular courses at Illinois!

# Part 0: Your Group!

Edit the next Python cell to add information about who you're working within your lab section:

```python
# First, meet your CAs and TA if you haven't already!
# ...first name is enough, we'll know who they are! :)
ta_name = ""
ca1_name = ""
ca2_name = ""


# Also, make sure to meet your team for this lab! Find out their name, what major they're in,
# and learn something new about them that you never knew before!
partner1_name = ""
partner1_netid = ""
partner1_major = ""
partner1_qotd_answer = ""

partner2_name = ""
partner2_netid = ""
partner2_major = ""
partner2_qotd_answer = ""

partner3_name = ""
partner3_netid = ""
partner3_major = ""
partner3_qotd_answer = ""

partner4_name = ""
partner4_netid = ""
partner4_major = ""
partner4_qotd_answer = ""
```

## Part 0.5 | Loading the Dataset

For the first part of our lab, we'll be using the GPA dataset we saw in class.

You can download the dataset by following to following link to the github storing the dataset: https://raw.githubusercontent.com/wadefagen/datasets/master/gpa/uiuc-gpa-dataset.csv

To download this file, you will want to **right click** and click **Save As**.

<!-- grade_id: cell-51fef9428e349da7, points: 0, grade: false, solution: true -->
```python

```

## Part 1 | Finding Total GPAs

Add two new columns, `Total Students` and `Average GPA`, that compute the total number of students in each course and the average weighted GPA for each course.

You should not include students with a `W` in either column and feel free to add other columns to help you out!

<!-- grade_id: cell-1043285927ac8a85, points: 0, grade: false, solution: true -->
```python
df["Total Students"] = ...
df["Average GPA"] = ...




```

<!-- grade_id: cell-46729cc74c213e52, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(df[
  (df["YearTerm"] == "2010-su") & \
  (df["Subject"] == "STAT") & \
  (df["Number"] == 410) \
]["Total Students"].sum() == 31), "It appears your \"Total Students\" column is incorrect."

assert(df[
  (df["YearTerm"] == "2018-fa") & \
  (df["Subject"] == "CS") & \
  (df["Number"] == 225) \
]["Total Students"].sum() == 697), "It appears your \"Total Students\" column is incorrect."

assert( abs(df[
  (df["YearTerm"] == "2018-fa") & \
  (df["Subject"] == "STAT") & \
  (df["Number"] == 100) \
]["Average GPA"].mean() - 3.247) < 0.01), "It appears your \"Average GPA\" column is incorrect."
```

## Part 2 | Finding the easiest and hardest courses

### Puzzle 2.1: Hardest Courses in the Dataset

Using functions you already know, find the **30 hardest courses** at Illinois based on their "Average GPA".

<!-- grade_id: cell-13f6d92b7a658526, points: 0, grade: false, solution: true -->
```python
df_hardestCourses = ...
df_hardestCourses


```

### Puzzle 2.2: Hardest Courses in 2019

Find the **30 hardest courses in 2019** at Illinois based on their "Average GPA".

<!-- grade_id: cell-a7106dca7694711c, points: 0, grade: false, solution: true -->
```python
df_hardestCourses_2019 = ...
df_hardestCourses_2019


```

### Puzzle 2.3: Easiest Courses in 2019

Find the **30 easiest courses in 2019** at Illinois based on their "Average GPA".

<!-- grade_id: cell-44d17a2ef57b1a19, points: 0, grade: false, solution: true -->
```python
df_easiestCourses_2019 = ...
df_easiestCourses_2019


```

<!-- grade_id: cell-af77ae707b9fa776, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error or output, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, RE-RUN your code, and then RE-RUN this cell.

assert( len(df_hardestCourses) == 30 ), "You should only have the 30 hardest courses."
assert( abs(df_hardestCourses["Average GPA"].mean() - 1.58) < 0.01 ), "Are you sure you have the 30 hardest courses?"

assert( len(df_hardestCourses_2019) == 30 ), "You should only have the 30 hardest courses in 2019."
assert( abs(df_hardestCourses_2019["Average GPA"].mean() - 2.04) < 0.01 ), "Are you sure you have the 30 hardest courses of 2019?"

assert( len(df_easiestCourses_2019) == 30 ), "You should only have the 30 easiest courses in 2019."
assert( abs(df_easiestCourses_2019["Average GPA"].mean() - 3.97) < 0.01 ), "Are you sure you have the 30 easiest courses of 2019?"
```

## Puzzle 3: Finding Data on  Every Subject

### Puzzle 3.1: Total Grades Given in Each Subject

Using `groupby` and `agg`, find the total number of each grade (eg: "A+"s, "A"s, "A-"s, etc.) for every subject in the DataFrame `df_grouped_subject`.

<!-- grade_id: cell-7d31cf501bfd8025, points: 0, grade: false, solution: true -->
```python



df_grouped_subject = ...
df_grouped_subject


```

<!-- grade_id: cell-e9ad64dffa5a1cf1, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3.1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, RE-RUN your code, and then RE-RUN this cell.

assert( len(df_grouped_subject) == 161 ), "Are you sure you grouped by subjects?"
assert( df_grouped_subject[ df_grouped_subject["Subject"] == "CS" ]["A"].sum() == 38472 ), "Are you sure you counted the right number of each subject?"
```

### Puzzle 3.2: Overall Average GPA

In `df_grouped_subject`, add a new column called `"Overall Average GPA"` that computes the total overall GPA of each subject at Illinois.

<!-- grade_id: cell-7d3923eba40e1f45, points: 0, grade: false, solution: true -->
```python
df_grouped_subject["Overall Average GPA"] = ...
df_grouped_subject


```

<!-- grade_id: cell-81b547cf0d42b67e, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3.2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error or output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, RE-RUN your code, and then RE-RUN this cell.

assert( abs(df_grouped_subject[ df_grouped_subject["Subject"] == "CS" ]["Overall Average GPA"].sum() - 3.249) < 0.01 ), "Your overall average GPA does not appear correct."
```

### Puzzle 3.3: Find the 15 easiest subjects at Illinois

Finally, find the 15 easiest subjects at Illinois based on their "Overall Average GPA"

<!-- grade_id: cell-cb713763b3781870, points: 0, grade: false, solution: true -->
```python
df_easiest_subjects = ...
df_easiest_subjects


```

### Puzzle 3.4: Find the hardest 15 subjects at Illinois

Find the hardest 15 subjects at Illinois based on their "Overall Average GPA"

<!-- grade_id: cell-eb5a80d62e1c7a45, points: 0, grade: false, solution: true -->
```python
df_hardest_subjects = ...
df_hardest_subjects


```

### Puzzle 3.5: Find your major!

Finally, find your major/subject in the DataSet and check out how hard it is! 

Feel free to look around and explore! Write your code below and write comments to explain what your code does. For example, what are the hardest classes in your major that have had at least 50 students?

<!-- grade_id: cell-116a2bb23ee81616, points: 0, grade: false, solution: true -->
```python

```

# Finale!

You're almost done -- congratulations!

You need to do two more things:

1. Save your work.  To do this, create a **notebook checkpoint** by using the menu within the notebook to go **File -> Save and Checkpoint**

2. After you have saved and checkpointed, exit this notebook, and return to the Data Science DISCOVERY page on how to use git to turn this notebook into the course!