# EC Notebook for Week #3: Working with columns

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

## 1. Loading pandas and importing data

For this notebook, we will use the "Hello" dataset -- as we did in lecture -- to answer a few questions on working with columns.

Using your @illinois.edu address, you can access the "Hello" dataset here: https://docs.google.com/spreadsheets/d/1uS6OoVsaCNEDwi-ozOXJe7Vn08M-41k7W1cdAlwLIiY/edit?usp=sharing

<!-- grade_id: cell-0d5f5f22a56cb05c, points: 0, grade: false, solution: true -->
```python

```

# 2. Do seniors students study less than freshman?

Using the "Hello" dataset, find the average number of hours freshman study (`freshman_avg_study_hours`) and the average number of hours seniors study (`senior_avg_study_hours`) each week:  

<!-- grade_id: cell-0c5e6b97fdf25f9a, points: 0, grade: false, solution: true -->
```python


freshman_avg_study_hours = ...
freshman_avg_study_hours


```

<!-- grade_id: cell-ba360e51c3559faa, points: 0, grade: false, solution: true -->
```python


senior_avg_study_hours = ...
senior_avg_study_hours


```

<!-- grade_id: cell-e03adddf03d4c81a, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(round(freshman_avg_study_hours ** 3, 2) == 31.53), "The average number of hours studying for freshman appears incorrect."
assert(round(senior_avg_study_hours ** 3, 2) == 47.26), "The average number of hours studying for seniors appears incorrect."
assert(round(freshman_avg_study_hours * senior_avg_study_hours, 2) == 11.42), "The average number of hours studying for seniors or freshman appears incorrect."


## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

# 3. Do introverts text less than extroverts?

In our dataset, you responded with the number of texts recieved yesterday.  Find the average number of texts sent by introverts and by extroverts:

<!-- grade_id: cell-0ded48531f57011e, points: 0, grade: false, solution: true -->
```python
introverts_avg_txts = ...
introverts_avg_txts


```

<!-- grade_id: cell-975d7137d0ca4494, points: 0, grade: false, solution: true -->
```python
extroverts_avg_txts = ...
extroverts_avg_txts


```

<!-- grade_id: cell-85b1350f5048af0d, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(introverts_avg_txts < extroverts_avg_txts), "This does not appear correct -- check your introvert and extrovert logic."


## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

## Submit Your Work!

Make sure to **Save and Checkpoint** your notebook, exit Jupyter, and submit your work! :)