# lab_simpsons-paradox

Hi Data Scientists,

Welcome to the third lab!  In this lab, you will explore using pandas to explore real-world datasets and find if Simpson's Paradox appears in several different datasets.


### Labs are Collaborative

In your lab section (right now!), find a person or two around you and form a group.  Work on this together, making sure all of you understand each component of the lab.


### Getting Help

Remember, there are a lot of ways to get help if you find yourself stuck:

1. In lab section, your TA, CAs, and your peers are here for you!


2. On the course Piazza page: https://piazza.com/illinois/spring2020/stat107


3. Office Hours:

  - Open office hours **every Monday, Tuesday, Wednesday, and Friday** from 4:00 pm - 6:00 pm in 23 Illini Hall

  - Prof. Wade + Prof. Karle are available every Wednesday from 8:30 am - 10:00 am in 2215 Siebel Center

# Puzzle 0: Your Group!

Edit the next Python cell to add information about who you're working within your lab section:

```python
# First, meet your CAs and TA if you haven't already!
# ...first name is enough, we'll know who they are! :)
ta_name = ""
ca1_name = ""
ca2_name = ""


# Also, make sure to meet your team for this lab! Find out their name, what major they're in,
# and learn something new about them that you never knew before!
partner1_name = ""
partner1_netid = ""
partner1_major = ""

partner2_name = ""
partner2_netid = ""
partner2_major = ""

partner3_name = ""
partner3_netid = ""
partner3_major = ""
```

# Puzzle 1: Homecoming Date Night

Imagine you and your partner are trying to find the perfect restaurant for a homecoming date night dinner!  Knowing this process can lead to hours of arguments, you seek out the answer to all of life’s questions: *online reviews*. 

The dataset `Black Dog Reviews.csv` (already provided for you as part of the lab code you checked out from git) contains data showing all of the males and females that have reviewed Black Dog in the past year and whether or not they recommended it.  The dataset `Papa Dels Reviews.csv` contains the same information for Papa Dels.


## 1.1: Data Cleaning

Unfortunately, this data was not as tidy as all of our data has been so far. :(

- Open `Black Dog Reviews.csv` and `Papa Dels Reviews.csv` CSV file using any program that can edit spreadsheets.  This can be Microsoft Excel (if you have it installed) or you can upload it to Google Drive and edit it with Google Sheets.
- Modify the file to ensure all rows that in the `Gender` column has the same value for the two gender categories (as `Male` and `Female`) found in this dataset.
- Finally, make sure to save the file as a CSV so you can used your cleaned dataset for this lab.


## 1.2. Loading Your DataSet

Load the DataSets you just cleaned into Python:

```python
# Import the pandas library as pd
import pandas as pd
```

<!-- grade_id: cell-3f3752d45165cffa, points: 0, grade: false, solution: true -->
```python
df_black_dog = ...
df_black_dog

```

<!-- grade_id: cell-1f967bc5dffe7c62, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1.2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.

assert(len(df_black_dog) == 40), "The dataframe does not appear to have correct dataset loaded"
## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

<!-- grade_id: cell-6171dbc0e3852e57, points: 0, grade: false, solution: true -->
```python
df_papa_del = ...
df_papa_del

```

<!-- grade_id: cell-630e1207191fb31e, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1.2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.

assert(len(df_papa_del) == 46), "The dataframe does not appear to have correct dataset loaded"
## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

## 1.3. Analysis

We want to use the Black Dog dataset to calculate the percentage of women who recommended Black Dog.  Then calculate the percentage of men who recommended Black Dog:

<!-- grade_id: cell-e8a1c3c40402d8e6, points: 0, grade: false, solution: true -->
```python
black_dog_female_percent = ...
black_dog_female_percent


```

<!-- grade_id: cell-fbdef0d3e0843326, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1.3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.

assert(black_dog_female_percent == 80), "The percentage of Female who recommended Black Dog does not appear to have been correctly calculated"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

<!-- grade_id: cell-e51a9d6c4f48495c, points: 0, grade: false, solution: true -->
```python
black_dog_male_percent = ...
black_dog_male_percent


```

<!-- grade_id: cell-00ebcd91c621f14a, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1.3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.

assert(black_dog_male_percent == 100*(1/3)), "The percentage of Male who recommended Black Dog does not appear to have been correctly calculated"
## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

Similarly, use the Papa Dels dataset to calculate the percentage of women who recommended Papa Dels.  Then calculate the percentage of men who recommended Papa Dels:

<!-- grade_id: cell-42037a70312202b7, points: 0, grade: false, solution: true -->
```python
papa_dels_female_percent = ...
papa_dels_female_percent


```

<!-- grade_id: cell-d534e36a406bf8da, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1.3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.

assert(papa_dels_female_percent == 90), "The percentage of Female who recommended Papa Dels does not appear to have correctly calculated"
## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

<!-- grade_id: cell-25efc5b635f23764, points: 0, grade: false, solution: true -->
```python
papa_dels_male_percent = ...
papa_dels_male_percent


```

<!-- grade_id: cell-b481131a6df56f62, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1.3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.

assert(papa_dels_male_percent == 50), "The percentage of Male who recommended Papa Dels does not appear to have been correctly calculated"
## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

### Observe the Results

Run the following cell to format all of your answers as a DataFrame:

```python
pd.DataFrame( [
    {'Female': black_dog_female_percent, 'Male': black_dog_male_percent },
    {'Female': papa_dels_female_percent, 'Male': papa_dels_male_percent },    
], index=['Black Dog', 'Papa Dels'])
```

❓ **Question** ❓: Should your homecoming date night be at Black Dog or Papa Dels?

- Does "Black Dog" (Row #1) have higher values or "Papa Dels" (Row #2) have higher values?

*Edit and replace this cell with your answer.*

## 1.4. Analysis - Part 2

For our final analysis, let's examine the **overall** percentage of people who recommend each restaurant:

<!-- grade_id: cell-76c29af59adb8ed8, points: 0, grade: false, solution: true -->
```python
black_dog_recommended_percent = ...
black_dog_recommended_percent

```

<!-- grade_id: cell-e9f12398cfd30bd5, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1.4 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.

assert(black_dog_recommended_percent == 62.5), "The percentage of Overall who recommended Black Dog does not appear to have been correctly calculated"
## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

<!-- grade_id: cell-e2051c256fd564b2, points: 0, grade: false, solution: true -->
```python
papa_dels_recommended_percent = ...
papa_dels_recommended_percent

```

<!-- grade_id: cell-b93794d7ed4b3f58, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1.4 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.

assert(papa_dels_recommended_percent == 100*(27/46)), "The percentage of Overall who recommended Papa Dels does not appear to have been correctly calculated"
## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

### Observe the Results

Run the following cell to format all of your answers as a DataFrame:

```python
pd.DataFrame( [
    {'Female': black_dog_female_percent, 'Male': black_dog_male_percent, 'Overall': black_dog_recommended_percent },
    {'Female': papa_dels_female_percent, 'Male': papa_dels_male_percent, 'Overall': papa_dels_recommended_percent },
], index=['Black Dog', 'Papa Dels'] )
```

❓ **Question** ❓: Should your homecoming date night be at Black Dog or Papa Dels?

- Does "Black Dog" (Row #1) have higher values or "Papa Dels" (Row #2) have higher values for **Overall**?

*Edit and replace this cell with your answer.*

# Puzzle 2: Shoe Size vs Shoes Owned

From the "Hello" dataset we created on the first day of class, we picked two random questions for us to explore:

- What is your shoe size?
- How many pairs of shoes do you have?

## 2.1: Load the Dataset

- You can find the "Hello" dataset here (make sure to log in with your @illinois.edu address to access it): https://docs.google.com/spreadsheets/d/1uS6OoVsaCNEDwi-ozOXJe7Vn08M-41k7W1cdAlwLIiY/edit#gid=0

<!-- grade_id: cell-a754c45cb34da497, points: 0, grade: false, solution: true -->
```python
# Read the hello dataset into a DataFrame
df_hello = ...
df_hello

```

<!-- grade_id: cell-96c94f328cfd0876, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2.1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(len(df_hello) == 91), "The dataframe does not appear to have correct dataset loaded"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

## 2.2: Observation Subsets

In this situation, let's define shoes that are strictly greater than size 9 to be "large" and shoes that less than or equal to size 9 are defined to be "small".

Create two DataFrames that contain subsets of our data: one DataFrame that includes everyone who has a "large" shoe size and one DataFrame that includes everyone who has a "small" shoe size:

<!-- grade_id: cell-8754c8095ea42040, points: 0, grade: false, solution: true -->
```python
df_large_size = ...
df_large_size

```

<!-- grade_id: cell-729252c9fbc53aed, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2.2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(len(df_large_size) == 47), "The dataframe does not appear to have been correct subset"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

<!-- grade_id: cell-b7ac40623aff2344, points: 0, grade: false, solution: true -->
```python
df_small_size = ...
df_small_size


```

<!-- grade_id: cell-3119e2c04a5d3f8d, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2.2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(len(df_small_size) == 44), "The dataframe does not appear to have been correct subset"
## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

## 2.3: Average Number of Pairs of Shoes

Find the average number of pairs of shoes of each group:

<!-- grade_id: cell-e4b369fb53d3613d, points: 0, grade: false, solution: true -->
```python
large_avg_shoes = ...
large_avg_shoes

```

<!-- grade_id: cell-303a6192d564894f, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2.3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(large_avg_shoes == 362/47),  "The average number of pairs of shoes in large shoe-size dataframe does not appear to seem correct. Make sure you are taking mean of the correct variable."
## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

<!-- grade_id: cell-d16479eedeb50e9c, points: 0, grade: false, solution: true -->
```python
small_avg_shoes = ...
small_avg_shoes

```

<!-- grade_id: cell-990a59a8df79bb00, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2.3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(small_avg_shoes == 503 / 44),  "The average number of pairs of shoes in small shoe-size dataframe does not appear to seem correct. Make sure you are taking mean of the correct variable."
## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

## 2.4. Observations

Answer the next two questions with 1-2 sentences each.  The cells after the question text are editable -- you should double-click and replace the text with your answer:

### Q1: Is there a difference in the number of shoes owned between the two groups?

*Edit and replace this cell with your answer.*

### Q2: What might be some possible confounding variable in this result?

*Edit and replace this cell with your answer.*

# Finale!

You're almost done -- congratulations!

You need to do two more things:

1. Save your work.  To do this, create a **notebook checkpoint** by using the menu within the notebook to go **File -> Save and Checkpoint**

2. After you have saved and checkpointed, exit this notebook and return to the Data Science Discovery page on how to use git to turn this notebook into the course!