# Part 0: Your Group!

Edit the next Python cell to add information about who you're working within your lab section:

```python
# First, meet your CAs and TA if you haven't already!
# ...first name is enough, we'll know who they are! :)
ta_name = ""
ca1_name = ""
ca2_name = ""


# Also, make sure to meet your team for this lab! Find out their name, what major they're in,
# and learn something new about them that you never knew before!
partner1_name = ""
partner1_major = ""
partner1_qotd_answer = ""

partner2_name = ""
partner2_major = ""
partner2_qotd_answer = ""

partner3_name = ""
partner3_major = ""
partner3_qotd_answer = ""

partner4_name = ""
partner4_major = ""
partner4_qotd_answer = ""
```

# Happy Birthday!

Just kidding, it's *(probably, with P=0.9973)* not your birthday, but it might be someone's birthday in this room! Taylor Swift says she will come and give a free concert on THE QUAD on March 14th if it is someone's birthday on that day. Now obviously the students want to make this happen, so they ask you to do some calculations regarding the probability of the concert happening.

- For these puzzles, we assume that each year is 365 days (ignore leap years, even though this year is one!)
- We also assume that everyone has an equal chance of a birthday each day *(not actually true IRL)*

## Review: Random Numbers in Python

Python provides an extremely easy way to generate a random integer (whole) number.  You need to take only two steps.

#### Step 1: Import Libraries

1\. `import` the `random` and `pandas` library

<!-- grade_id: cell-18f4035da8f5ab1e, points: 0, grade: false, solution: true -->
```python

```

#### Step 2: Random Number

2\. Use the function `random.randint(x, y)`, where `x` and `y` are numbers, which will return in the inclusive range `x` to `y`.

- Example: `random.randint(4, 10)` randomly chooses `4`, `5`, `6`, `7`, `8`, `9` or `10`
- Example: `random.randint(0, 1)` randomly chooses `0` or `1`

In the following cell, generate a random number that represents a day in the year.  (**HINT**: *How many days are there every year?*)

<!-- grade_id: cell-c2dd17defd9fa8a2, points: 0, grade: false, solution: true -->
```python
some_day = ...


```

## Puzzle 1: My Birthday

The **birthday problem** is a classic problem we covered in lecture that asks: *How likely is it for at least two people in a room of n people to share a birthday?*  In other words, what's the chance that there's at least one match of birthdays in a room with n people? We'll cover this Puzzle \#2.


### "My Birthday Problem"

Instead of doing the same thing in lecture, let's consider **your birthday**: 

> *How likely is it for at least one person in a room to have the same birthday as **me**?*

### One Other Person (n=1)

If you are in a room with **one other person**, what is the probability you both have the same birthday? (*Of course you know your own birthday!*) There is no need for any python code for this puzzle, it should be a mathematical answer!

<!-- grade_id: cell-16021a893dda10f8, points: 0, grade: false, solution: true -->
```python
p_one_person_my_birthday = ...
p_one_person_my_birthday


```

### Two Other People (n=2)

If you are in a room with **two other people**, what is the probability that at least one of them has the same birthday as you?

<!-- grade_id: cell-d776448804b67dfc, points: 0, grade: false, solution: true -->
```python
p_two_others_my_birthday = ...
p_two_others_my_birthday


```

### Three Other People (n=3)

If you are in a room with **three other people**, what is the probability that at least one of them has the same birthday as you?

<!-- grade_id: cell-9b6271534fe56ab1, points: 0, grade: false, solution: true -->
```python
p_three_others_my_birthday = ...
p_three_others_my_birthday


```

<!-- grade_id: cell-2213109ad4022435, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1a ==
# - This read-only cell contains test cases for your previous cells.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(abs(p_one_person_my_birthday - 0.0027397) < 0.001), "The probability for one birthday doesn't seem right."
assert(abs(p_two_others_my_birthday - 0.0054719) < 0.001), "The probability for two birthdays doesn't seem right."
assert(abs(p_three_others_my_birthday - 0.0081966) < 0.001), "The probability for three birthdays doesn't seem right."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

### `n` Other People

Suppose we have a room with **`n` other people**.  Complete the function `p_my_birthday(n)` that returns the probability that at least one person has the same birthday as you:


<!-- grade_id: cell-9fcc75505fada2fa, points: 0, grade: false, solution: true -->
```python
def p_my_birthday(n):
    ...
    

```

### Running Your Function:

Here are a few functions to test our function.  Make sure the answers are what you expect. You do not need to do anything for this cell, just run it!

```python
print('A room with one other person:', p_my_birthday(1))
print('A room with two other people:', p_my_birthday(2))
print('A room with three other people:', p_my_birthday(3))
print('A room with 100 other people:', p_my_birthday(100))
print('A room with 1000 other people:', p_my_birthday(1000))
print('A room with 10000 other people:', p_my_birthday(10000))
```

<!-- grade_id: cell-f1a2b09260950b38, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1b ==
# - This read-only cell contains test cases for your previous cells.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(abs(p_my_birthday(1) - 0.0027397) < 0.001), "The probability doesn't seem right."
assert(abs(p_my_birthday(2) - 0.0054719) < 0.001), "The probability doesn't seem right."
assert(abs(p_my_birthday(3) - 0.0081966) < 0.001), "The probability doesn't seem right."
assert(abs(p_my_birthday(100) - 0.23993292) < 0.001), "The probability doesn't seem right."
assert(abs(p_my_birthday(1000) - 0.93565) < 0.001), "The probability doesn't seem right."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

## Result Plot

The following code below uses your `p_my_birthday` function to generate a plot showing the probability of you having the same birthday as someone else increasing as the number of people in the room grows:

```python
# Create a new empty list, append 2,000 results to it:
results = []
for i in range(2000):
    results.append( {'Other People': i, 'P(at least one person shares my birthday)': p_my_birthday(i) } )
    
# Create a DataFrame out of that list:
df = pd.DataFrame( results )

# Show a plot of our DataFrame:
plot = df.plot.line(y='P(at least one person shares my birthday)', x='Other People', grid=True)
#
```

## Puzzle 2: The Classic Birthday Problem

The **birthday problem** is a classic problem that asks: *How likely is it for at least two people in a room to share a birthday?*


### A One-Person Party (n=1)

Let's start with just one person.  The probability that they share a birthday with someone else is... well... 0%. :(

$P(shared\ birthday\ with\ 1\ person) = 0$


### A Two-Person Party (n=2)

The probability of two people sharing the same birthday is non-zero!  This is the probability of the second person having exactly the same birthday as the first -- 1 out of 365!

$P(shared\ birthday\ with\ 2\ people) = \frac{1}{365}$


### A Three-Person Party (n=3)

Now it gets interesting!  In a three person party, what is the probability that there is **at least one** match of birthdays?

Use the following cell to find the answer in Python, the solutions will be mathematical equations:

<!-- grade_id: cell-9db7fd67d49459bd, points: 0, grade: false, solution: true -->
```python
P_three_person_party = ...
P_three_person_party


```

<!-- grade_id: cell-3aa5b5a4afef9fbc, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2 (with n=3) ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(abs(P_three_person_party - 0.0082041) < 0.001), "The probability doesn't seem right."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

### A Four-Person Party (n=4)

In a four person party, what is the probability that there is **at least one** match of birthdays?

Use the following cell to find the answer in Python:

<!-- grade_id: cell-1a5b32da97b97164, points: 0, grade: false, solution: true -->
```python
P_four_person_party = ...
P_four_person_party



```

<!-- grade_id: cell-571231fcdfa9de08, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2 (with n=4) ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(abs(P_four_person_party - 0.01635) < 0.001), "The probability doesn't seem right."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

### An `n`-person Party

Create a function `P_shared_birthday` that returns the probability that there is **at least one** match of birthdays?

<!-- grade_id: cell-8826e853f191d14f, points: 0, grade: false, solution: true -->
```python
def P_shared_birthday(n):
    ...
       



# Here's a few tests that will 
print('A room with just you:',  P_shared_birthday(1))
print('A room with two people:',  P_shared_birthday(2))
print('A room with three person:',  P_shared_birthday(3))
print('A room with four person:',  P_shared_birthday(4))
print('A room with 40 other person:',  P_shared_birthday(40))
print('A room with 80 other person:',  P_shared_birthday(80))
print('A room with 400 other person:',  P_shared_birthday(400))
```

<!-- grade_id: cell-5a323ce5fab0c281, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2 (function) ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(abs(P_shared_birthday(3) - 0.0082041) < 0.001), "The probability doesn't seem right."
assert(abs(P_shared_birthday(4) - 0.01635) < 0.001), "The probability doesn't seem right."
assert(abs(P_shared_birthday(40) - 0.8912318) < 0.001), "The probability doesn't seem right."
assert(abs(P_shared_birthday(80) - 0.999914) < 0.001), "The probability doesn't seem right."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

## Result Plot

Create a plot showing the growth of `P_shared_birthday` similar to the plot generated in Puzzle \#1.  *Make sure the range of your graph focuses on the interesting features.* (Hint: You will have to make a data frame using the data generated from your 'P_shared_birthday' function.)

Take a look at the documentation for the plot you will be using: https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.plot.line.html

<!-- grade_id: cell-c46ee253ad4a8777, points: 0, grade: false, solution: true -->
```python
df = ...


```

## Puzzle 3: Taylor's Birthday Problem

Taylor Swift wants **every day to be someone's birthday**!  Given a massive concert of `n` people, what is the probability that **all 365 days have at least one person with a birthday that day**?

Since Wade, Karle, and the full course staff <3s Taylor, we wrote a function to do this simulation for Taylor!  You can check out the code below, but all you need to know is that:

> The `simulation_for_taylor` function returns one simulation of the number of people needed to satisfy Taylor's wish for there to be someone with a birthday for every day of the year at her concert!

...you will need to **use this function** to help us calculate probabilities!

```python
# The `simulation_for_taylor` function returns one simulation of the number of people needed
# to satisfy Taylor's wish for there to be someone with a birthday for every day of the year
# at her concert!
#
# @return: One simulation of the number of people needed at the concert.
#   Examples: 2056, 2520, 2868, 3033, 2022, 1743, ...

def simulation_for_taylor():
    ct = 0   #< Count the number of people needed
    day = 0  #< Count the current unique days found
    
    # While all 365 unique days have not been seen:
    while day < 365:
        # Add a day:
        ct = ct + 1
        
        # Check if a random number is a new, unique day.
        # - When day == 0, we've seen no unique days so P(unique day) == 1 as randint(0, 364) will always be `>= 0` (100%)
        # - When day == 1, P(unique day) == (364/365) and we check if we get a number `>= 1` (any number but 0)
        # - When day == 2, P(unique day) == (363/365) and we check if we get a number `>= 2` (anything but 0 or 1)
        # - ...
        # - When day == 364, P(unique day) == (1/365) so we need get randint(0, 364) to be exactly 364.
        if random.randint(0, 364) >= day:
            day = day + 1
            
    # Return the number of people needed
    return ct    
```

### Puzzle 3.1: Run the already-complete simulation

Run this next cell several times.  See the different number of people at Taylor's concert for everyone to have a birthday:

<!-- grade_id: cell-a60a84fa1d64c87d, points: 0, grade: false, solution: true -->
```python

```

### Puzzle 3.2: Collection Simulation Results

Using `simulation_for_taylor` as the simulation, collect the results of running the simulation at least **1000** times and store the result in `df_taylor`:

*(Note: This may take a few seconds or up to a minute to run on a slower laptop.  You may want to just run the simulation 10 times until your code works before changing to to 1,000.)*

<!-- grade_id: cell-6a44b9f35daf8325, points: 0, grade: false, solution: true -->
```python
df_taylor = ...



df_taylor.sample(10)
```

### Puzzle 3.3: Finding the probability we delivered for Taylor

The concert is held and we have **exactly 2,000 people packed onto the quad for Taylor Swift!**  Using your simulation results, find an estimate for the probability that Taylor will be happy to know every day of the year is someone’s birthday:

<!-- grade_id: cell-0b48a78d683ce577, points: 0, grade: false, solution: true -->
```python
# First, find the times the simulation was a success:
df_success = ...



df_success.sample(5)
```

<!-- grade_id: cell-3cd0d620430a89ba, points: 0, grade: false, solution: true -->
```python
P_everyBirthday = ...
P_everyBirthday


```

### Puzzle 3.4: Distribution Histogram

Create a histogram of your `df_taylor` results, showing the distribution of people needed in our simulation to account for every birthday.

<!-- grade_id: cell-62147f75639a94d3, points: 0, grade: false, solution: true -->
```python
# Create a histogram:



```

<!-- grade_id: cell-1bae19913f8df080, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert( len(df_taylor) >= 1000 ), "Ensure you have done at least 1,000 simulations."
assert( len(df_success) > 0 ), "Ensure your success includes all simulations with 2,000 or fewer people."
assert( len(df_success) < len(df_taylor) ), "Ensure your success includes all simulations with 2,000 or fewer people."
assert( P_everyBirthday > 0.1 and P_everyBirthday < 1 ), "Ensure `P_everyBirthday` is a probability."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()
print("Great! You figured it out. Now go start inviting people!")
```

## Submit Your Work!

Make sure to **Save and Checkpoint** your notebook, exit Jupyter, and submit your work! :)