# EC Notebook for Week 2: Blocking and Conditionals

This extra credit Python notebook will let you practice the material you saw in lecture!  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107!

This notebook is worth +1 if turned in before 11:59pm on **Friday, Feb. 7**.

## 1. Loading pandas and importing data

For this notebook, we will continue to use the "Hello" dataset introduced last lecture!

Using your @illinois.edu address, download the dataset as a CSV and **make sure to save or move it to the same directory as this notebook**: https://docs.google.com/spreadsheets/d/1uS6OoVsaCNEDwi-ozOXJe7Vn08M-41k7W1cdAlwLIiY/edit?usp=sharing

Use the following cell to load the pandas library and read the DataSet as a DataFrame:

<!-- grade_id: cell-7543f21c9c406c26, points: 0, grade: false, solution: true -->
```python
...

```

## 2. Creating a Control Group

In lecture, we worked out a puzzle.  From the lecture handout:

> Use Python to create a control group of 20 students, using blocking to control for the influence we suspect having more than 1 siblings might have.  Make sure you have an equal number of students with 0 or 1 sibling and an equal number of students with 2 or more siblings.

We will complete the code in a notebook this in four separate steps to get practice on each part.

### 2.1: Sibling Block #1 -- Zero or One Siblings

Creating a DataFrame that contains all students with 0 or 1 siblings and call it `block1`:

<!-- grade_id: cell-42b76d09e5d1ea90, points: 0, grade: false, solution: true -->
```python
block1 = ...



# Show a few rows to check if the block has the data we expect:
block1.sample(n=3)
```

### 2.2: Sibling Block #2 -- Two or More Siblings

Create a DataFrame that contains all students with 2 or more siblings and call it `block2`:

<!-- grade_id: cell-ad3eef9bc1500907, points: 0, grade: false, solution: true -->
```python
block2 = ...


# Show a few rows to check if the block has the data we expect:
block2.sample(n=3)
```

### 2.3: Selecting equal sized groups from block1 and block2

**Randomly** select an equal number of students from each of `block1` and `block2` and store them in `control1` and `control2`.  We will combine both groups to make the full control group:

<!-- grade_id: cell-71025862479973ff, points: 0, grade: false, solution: true -->
```python
control1 = ...


# Show a few rows to check if the control group has the data we expect:
control1.sample(n=3)
```

<!-- grade_id: cell-0a28afaa87e1114d, points: 0, grade: false, solution: true -->
```python
control2 = ...



# Show a few rows to check if the control group has the data we expect:
control2.sample(n=3)
```

#### 2.4. [CHALLENGE!] -- Combine the two sub-groups into a single control group

Use the following code to merge the contents of `control1` and `control2` together into one large group called `control`.  To do this, you will need to use the `pd.concat` command that you can **find on your cheat sheet**.

<!-- grade_id: cell-f5b1dce206b67c5a, points: 0, grade: false, solution: true -->
```python
control = ...



# Show the full control group (it should have 20 rows, mixed number of siblings)
control
```

<!-- grade_id: cell-78cde41a46e52f67, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2 ==
# - This read-only cell contains test cases for your previous cells.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert( len(control) == 20 ), "Your control group must only have 20 courses in it."
assert( len(control[control.Siblings < 2] ) == 10 ), "Your control group must only have 10 people with 0-1 siblings in it."
assert( len(control[control.Siblings >= 2] ) == 10 ), "Your control group must only have 10 people with 2+ siblings in it."


## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:

tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

## Submit Your Work!

Make sure to **Save and Checkpoint** your notebook, exit Jupyter, and submit your work! :)