# Find your Data Science DISCOVERY twin!

Some of you may know your tablemates in this class from earlier in the semester. However, have you ever wondered if there is someone even more similar to you in the class than the people you already know? In this lab you will find your in-class doppelganger. While COVID-19 might separate us, you will learn more about your peers than ever before through the power of Data Science!

In this lab, you will explore data frames more using Pandas again! As you go through the lab:

- Make sure to run every cell with Python code and observe the output
- Complete all puzzles and submit this lab before Monday, May 4th at 11:59pm for 105% credit (or Wednesday, May 6th at 11:59pm for 100% credit)

Remember, there are several ways to get help:
- Post on Piazza
- Attend any of the Zoom-based Office Hours

You need to look at the Day 1 Dataset (from the first week of class). Here's the link (make sure to login with @illinois): https://docs.google.com/spreadsheets/d/1uS6OoVsaCNEDwi-ozOXJe7Vn08M-41k7W1cdAlwLIiY/edit?usp=sharing


Have fun!

## Part 0 | Loading the Dataset

For the first part of our lab, we'll be using the "Hello" dataset one last time.

- Import the proper libraries
- Convert NaN values to 0 using `df.fillna(0)`
- Output a description of the data frame

<!-- grade_id: cell-193e1aa1a2845076, points: 0, grade: false, solution: true -->
```python
# Always first, libraries:
import seaborn as sns
...


```

<!-- grade_id: cell-b13bbec99d775cf3, points: 0, grade: false, solution: true -->
```python
# Load your CSV:
df = ...


# Convert NaN values to 0 using `df.fillna(0)`:
df = df.fillna(0)

# Output a few rows:
df.sample(n=3)
```

<!-- grade_id: cell-4630e717d7a78a09, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 0 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs with the success message as your output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(len(df) == 91), "The dataframe does not appear to have correct dataset loaded"
assert(df.isnull().values.any() == False), "The dataframe still seems to have NaN values, make sure to replace those with 0"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

## Part 1 | Numerical vs. Categorical Data

**Numerical (or Quantitative) Data** is any type of data that can be expressed through numbers instead of descriptions. This is how mathematical operations can be applied to data.

**Categorical (or Qualitative) Data** is any type of data that can be placed into categories using names or groups. Some categorical variables can take numerical values (like zip codes or your social security number).
*A good trick to telling the difference between those is to ask yourself if it would make sense to sum two of the values. For example, the sum of two zipcodes is not meaningful.*


### Puzzle 1.1a: Getting the Numerical Data

Here we give you a function with some initial documentation to find all the columns from a dataframe that are numeric. 

- Use the given function to return a list of numeric columns and store that in the variable `numeric`
- Once you've made your list named `numeric` use the code given to make a dataframe of just numeric data.

**There is no need to edit the function code or the code creating `df_numeric`**

<!-- grade_id: cell-cb7ad7b8f8f90b4e, points: 0, grade: false, solution: true -->
```python
# isNumeric Function Description:
#     input: an arbitrary DataFrame
#     tests to see if a column is numeric
#     return: a list of the the numeric columns
def isNumeric(df):
    from pandas.api.types import is_numeric_dtype

    numericColumns = []
    for column in df:
        if is_numeric_dtype(df[column]):
            numericColumns.append(column)
    return numericColumns

# Use the given function to return a list of numeric columns and store that in the variable `numeric`:
numeric = ...
numeric


```

```python
# Create a DataFrame out of only columns contained in `numeric`
# (This code is provided and already complete! :))
df_numeric = df[ numeric ]
df_numeric.sample(n=5)
```

<!-- grade_id: cell-48f03cdf8ec4691b, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1.1a ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs with the success message as your output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(len(numeric) == 22), "The list of numeric columns seems to have an incorrect number of elements"
assert(df_numeric.shape[0] == 91), "The dataframe doesn't have the right number of rows"
assert(df_numeric.shape[1] == 22), "The dataframe doesn't have the right number of columns"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

### Puzzle 1.1b: Getting the Categorical Data

This time you're going to create a function to find which columns from a dataframe are categorical. Model this function after `isNumeric` but instead of using `is_numeric_dtype()` to test if the column data is numerical, use `is_string_dtype()` to test if the column has text data.

- Make a function that finds a list of categorical columns
- Store those columns in the variable `categorical`
- Create a dataframe of just the categorical variables called `df_categorical`

<!-- grade_id: cell-7f8aafe93eb47802, points: 0, grade: false, solution: true -->
```python
# isCategorical Function Description:
#     input: an arbitrary DataFrame
#     tests to see if a column is categorical
#     return: a list of the the categorical columns

def isCategorical(df):
    from pandas.api.types import is_string_dtype
    
    # Complete this function.
    # - Model this function after isNumeric but instead of using is_numeric_dtype()
    #   to test if the column data is numerical, use is_string_dtype() to test if
    #   the column has text data.
    
    
    
    return ...

categorical = ...
categorical


```

```python
df_categorical = df[ categorical ]
df_categorical.sample(n=5)
```

<!-- grade_id: cell-6f2c69f64d958645, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1.1b ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs with the success message as your output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(len(categorical) == 11), "The list of numeric columns seems to have an incorrect number of elements"
assert(df_categorical.shape[0] == 91), "The dataframe doesn't have the right number of rows"
assert(df_categorical.shape[1] == 11), "The dataframe doesn't have the right number of columns"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

### Puzzle 1b) Visualizing Numerical Data

Plot the histograms for each columns in `df_numeric`

- If needed: In `df.hist`, use **`figsize=(30, 40)`** as an extra argument passed to `hist` to make your figure bigger.

<!-- grade_id: cell-320d0db0b0449486, points: 0, grade: false, solution: true -->
```python
#plot the histogram for each column in df_numeric


```

## Part 1c) Visualizing Categorical Data

To visualize the categorical data we will be using a bar plot.  However, we first need to convert our DataFrame with the categorical data into counts for each value.  To do this, we can use [`Series.value_counts()`](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.Series.value_counts.html) on a column of data.

Let's see how this works.  Select just the `["School Year"]` from `df_categorical` and apply `.value_counts()`.  Then, in the next cell, use that data to make a bar plot using `.plot.bar()`:

<!-- grade_id: cell-6a70b687d4464dd2, points: 0, grade: false, solution: true -->
```python
# Use `.value_counts()` to find the values counts of data about ["School Year"]
value_counts = ...
value_counts


```

<!-- grade_id: cell-9a11d63567013a46, points: 0, grade: false, solution: true -->
```python
# Now, use `value_counts` to plot a bar chart:



```

### Repeat for all categorical data: 

Finally, repeat this process for all of the columns in `df_categorical`:
- "School Year"
- "Phone"
- "Computer"
- "Hot Dog is Sandwitch"
- "Programmed before?"
- "Personality"
- "Pineapple on Pizza"

```python
for column in ["Phone", "Computer", "Hot Dog is Sandwitch", "Programmed before?", "Personality", "Pineapple on Pizza"]:
    import matplotlib. pyplot as plt
    df_categorical[column].value_counts().plot.bar(title=column)
    plt.show()
```

## Part 3 | Now, Time to Find Yourself

You cannot find people you are compatible with before finding yourself... literally! 

### Puzzle 3.1: Find Yourself

In the following cell, write the Python code to create a DataFrame with only yourself:

<!-- grade_id: cell-04bc136da68bc695, points: 0, grade: false, solution: true -->
```python
myrow = ...
myrow


```

<!-- grade_id: cell-7b019251e3c71d1f, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3.1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs with the success message as your output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(len(myrow) == 1), "The row of yourself is not correctly set"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

### DataFrame with one row vs. a Series

In `pandas`, there are three major data structures we will be using:

- A **DataFrame** is similar to an Excel sheet.  It contains zero or more rows and zero or more columns.  It can contain up to an unlimited number of rows and columns and is the basic data structure we will work with in Data Science Discovery

- A **GroupBy** is a grouping of rows within a DataFrame.  This "three dimensional data" requires us to use `agg` to reduce it back down to a **DataFrame**.

- A **Series** is a single row with zero or more values (columns).  You can think of a DataFrame as a collection of **Series** objects.  When we deal with a single row, we will usually want to deal with a **Series** instead of a **one row DataFrame**.

#### Converting between types

The `pandas` library makes it easy to convert between a **Series**, **DataFrame**, and **GroupBy**.

![image.png](img/datatypes.png)

Expressed as a list:

- From a **Series** (`row`) to a **DataFrame**
  * `row.to_frame()`
- From a **DataFrame** (`df`) to a **GroupBy**:
  * `df.groupby(by='columnName')` *-- (You know this already!)*

- From a **GroupBy** (`group`) to a **DataFrame**:
  * `df.agg( function )` *-- (You know this already!)*
- From a **DataFrame** (`df`) to a **Series**:
  * `df.squeeze()`, if the DataFrame has one row


  

### Puzzle 3.2: Transform your row into a Series
 
Using the correct function from above, transform the row containing only you from a one-row DataFrame into a Series.  Make sure the Series is named `myrow` (this will match up with code we provide you with later):

<!-- grade_id: cell-8c946526ed7a65cd, points: 0, grade: false, solution: true -->
```python
# Your code here:
...


```

<!-- grade_id: cell-13e6051988c8bf2d, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3.2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs with the success message as your output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(type(myrow) == pd.core.series.Series), "You did not successfully transform myrow's type into series"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

## Part 4 | Create a similarity function

The first step to finding who is the most similar to you is to create a **function** that will generate a similarity value for each person (or row). Think about the **algorithm** you would do this without code:

1. Go through each person.
2. Check how many questions that row has answered the same thing as you did.
3. Record the number of matches in that row as a "similarity score".

### Creating a Python function

Unlike `sum`, `count`, or other functions, Python does not provide us a **function** that finds the similarity between two rows.  Instead, we need to write it ourselves. 😞

Using what you learned in lecture, let's write this function!

**Helpful Tips:**
- Remember that you can loop through the columns of a dataframe
- Treat numeric data and categorical data differently
- Have a plan to test for if a column is all 0's

For a measure of ***closeness*** between values for the numerical data we'll use the following formula:

$1 -$ ( $|actual - desired|$ / $(maxColumnValue - minColumnValue)$ )


<!-- grade_id: cell-51500f7d58876da9, points: 0, grade: false, solution: true -->
```python
def similarity(row):
    score = 0
    
    for column in df:
        ...
    return score

# First test this function by checking how similar you are with yourself.
# The expected result will be a number equal to the number of columns you checked:


```

## Part 5 | Run the similarity function on all rows

Now, finally, DATA SCIENCE!

Add a column to your original DataFrame that contains the value of `df.apply(similarity)` using the proper axis which will run your `similarity` function you defined earlier across all rows:

<!-- grade_id: cell-8ddb7d2fa7a20c87, points: 0, grade: false, solution: true -->
```python
#your code to add a column to dataframe called "similarity"


```

<!-- grade_id: cell-ec38b4dade899048, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 5 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs with the success message as your output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(df.shape[1] >= 34), "You need to apply your similarity function to your dataframe"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

## Part 6 | Find your twin!

Your last column in your DataFrame is a **metric** (a measurement) of how similar you are to everyone else in the class!

Use your **cheat sheet** (link: https://pandas.pydata.org/Pandas_Cheat_Sheet.pdf) to find how to sort a DataFrame where the most similar person is on top.  Store that sorted DataFrame in `df_twins`.

- You should be most similar to yourself (you will always match yourself)!  *If this is not the case, you did something wrong and should double-check all of your steps up to this point.*

<!-- grade_id: cell-2f540c719178869e, points: 0, grade: false, solution: true -->
```python
df_twins = ...
df_twins


```

<!-- grade_id: cell-d17c8bf547ab9ea7, points: 3, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 6 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs with the success message as your output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert( df_twins.head(1).squeeze()["Name"] == myrow["Name"] ), "You must be most similar to yourself."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

# Finale!
## Submit Your Work!

You're almost done -- congratulations!

You need to do two more things:

1. Save your work.  To do this, create a **notebook checkpoint** by using the menu within the notebook to go **File -> Save and Checkpoint**

2. Choose `File` and then `Close and Halt` from this notebook.

3. Choose `Quit` on the main notebook webpage.

4. Return to your command line and follow the directions on the Data Science Discovery page on how to use git to turn this notebook into the course!