# Welcome to lab k-means(karle)!

We will be exploring k-means clustering! 
I promise this won't be scary :)

Really, loops are an incredibly useful tool that will help you code for other purposes in life. We're going to start by introducing you to while and for loops and then we will apply them to clustering.

## Part 1: Import the libraries and data sets you need

### Libraries

To complete this lab, you will need two libraries:

1. `pandas`, which I **know** you know how to import :)
2. The `KMeans` model, which can be imported with `from sklearn.cluster import KMeans`

### Datasets

Additionally, you be using the beer dataset.  Here is the direct link to the CSV dataset: http://courses.las.illinois.edu/spring2020/stat107/static/data/beer-dataset.csv

Make sure to load the dataset into Python! :)

<!-- grade_id: cell-47bb214090f380cf, points: 0, grade: false, solution: true -->
```python
# Import libraries:


# Load dataset:
df = ...
df


```

## Part 1.1: EDA - Let's do some statistics! :)

Before going any further, let's find out a few basic statistics about the dataset.

- Use `df.describe()` to find basic EDA about the dataset.  Store the EDA table into `df_eda` below.
- Use `df.corr()` to find the correlation between data in the dataset.  Store the correlation table in `df_corr` below.

<!-- grade_id: cell-322e4f6a4d5584e8, points: 0, grade: false, solution: true -->
```python
df_eda = ...
df_eda


```

<!-- grade_id: cell-10a7ec62ba539341, points: 0, grade: false, solution: true -->
```python
df_corr = ...
df_corr


```

<!-- grade_id: cell-80b7b810a8e3351f, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert('Carbohydrates' in df_eda), "Make sure you've done EDA"
assert('Carbohydrates' in df_corr), "Make sure you've done correlation"

print("Congrats!🎉😊 All tests passed!")
```

## Part 2: Beer Dataset Clustering!

Similar to what you did in lecture this week, we are going to be looking at the dataset Karle found! The dataset has these five variables:

- Brand
- Brewery
- PercentAlcohol
- Calories
- Carbohydrates

In lecture you found the regression equation for predicting Calories from PercentAlcohol and Carbohydrates. Now, we are going to cluster the dataset using `KMeans`.

### Part 2a: Create and Fit a KMeans model

1. Create the model with **two** clusters
2. Fit your model with your data

Each of these two steps are just one line of code, for two lines of code below:

<!-- grade_id: cell-1572792b891915cc, points: 0, grade: false, solution: true -->
```python
model = ...
...


```

### Part 2b: Label data with a cluster

Use the model you created a fit to create a new column that contains the predicted cluster:

<!-- grade_id: cell-54a882b2799f7a7c, points: 0, grade: false, solution: true -->
```python
df['cluster'] = ...
df


```

### Part 2c: Explore the Clusters

Find out what beers appear in `cluster` `1`.  Store those in df_cluster1:

<!-- grade_id: cell-c6eb2274440818ee, points: 0, grade: false, solution: true -->
```python
df_cluster1 = ...
df_cluster1


```

<!-- grade_id: cell-7d79328221c654ff, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert('cluster' in df), 'There is no column `cluster` in `df`.'
assert( len(df[ df['cluster'] < 2 ]) == 85), 'Make sure your `KMeans` is finding only two clusters.'
assert( len(df_cluster1) == 26), 'Make sure your `KMeans` is finding only two clusters.'

print("Congrats!🎉😊 All tests passed!")
```

## Part 3: Iris Flowers

*You're doing awesome!  Let's try this with a completely new dataset you haven't seen yet!*

One of the **most classic** datasets in statistics is the "Iris Flower dataset".  This dataset is a collection of five variables:

- Sepal Length (cm)
- Sepal Width (cm)
- Petal Length (cm)
- Petal Width (cm)
- Species Classification


In the dataset, there are three species present:

![iris.jpg](iris.jpg)

This data is so classic that it is included as part of the `sklearn` Python library we have been using to do various forms of machine learning so far!  The following cell will load the dataset into a DataFrame `df`:

```python
from sklearn import datasets
import pandas as pd

iris = datasets.load_iris()
df = pd.DataFrame(iris['data'], columns=iris['feature_names'])
df['species'] = iris['target']
df['species name'] = iris['target_names'][ iris['target'] ]
df.sample(5)
```

### Puzzle 3.1: Fit the Model

Build a k-means clustering model with **three** clusters and train the model using all four input measurements:

- Sepal Length (cm)
- Sepal Width (cm)
- Petal Length (cm)
- Petal Width (cm)

<!-- grade_id: cell-3d0e0ea42d5c1f6d, points: 0, grade: false, solution: true -->
```python
model = ...
...


```

### Puzzle 3.2: Predict the Species with k-means clustering

Use the k-means clustering model to predict the species, adding the predicted species as a new column called `predicted species`.

<!-- grade_id: cell-b7579689f52d26c9, points: 1, grade: true, solution: true -->
```python
df['predicted species'] = ...



df.sample(3)
```

### Puzzle 3.3: Find the Accuracy

Finally, find the percentage of observations where the predicted species is a **correct match** the actual species.  *(When is the `predicted species` equal to the actual `species`?)*

Find the percentage of times our predicted species is correct and store that percentage in `accuracy`:

<!-- grade_id: cell-5fb2e528a53e01e2, points: 0, grade: false, solution: true -->
```python
accuracy = ...
accuracy


```

<!-- grade_id: cell-b05cd80356703d30, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert('predicted species' in df), "Make sure you have a column named `predicted species`"
assert(accuracy > 80), "Your accuracy should be higher than what you calculated!"

print("Congrats!🎉😊 All tests passed!")
```

## Submit Your Work!

You're almost done -- congratulations!

You need to do two more things:

1. Save your work.  To do this, create a **notebook checkpoint** by using the menu within the notebook to go **File -> Save and Checkpoint**

2. Choose `File` and then `Close and Halt` from this notebook.

3. Choose `Quit` on the main notebook webpage.

4. Return to your command line and follow the directions on the Data Science Discovery page on how to use git to turn this notebook into the course!

```python

```